Comparative Brain-Wide Mapping of Isoflurane and Ketamine-Activated Nuclei and Functional Networks

Author Response

The following is the authors’ response to the previous reviews.

We appreciate the reviewers for their insightful feedback, which has substantially improved our manuscript. Following the suggestions of the reviewers, we have undertaken the following major revisions:

a. Concerning data transformation, we have adjusted the methodology in Figures 2 and 3. Instead of normalizing c-Fos density to the whole brain c-Fos density as initially described, we now normalize to the c-Fos density of the corresponding brain region in the control group. b. We have substituted the PCA approach with hierarchical clustering in Figures 2 and 3.

c. In the discussion section, we added a subsection on study limitations, focusing on the variations in drug administration routes and anesthesia depth.

Enclosed are our detailed responses to each of the reviewer's comments.

Reviewer #1:

1a. The addition of the EEG/EMG is useful, however, this information is not discussed. For instance, there are differences in EEG/EMG between the two groups (only Ket significantly increased delta/theta power, and only ISO decreased EMG power). These results should be discussed as well as the limitation of not having physiological measures of anesthesia to control for the anesthesia depth.

1b. The possibility that the differences in fos observed may be due to the doses used should be discussed.

1c. The possibility that the differences in fos observed may be due kinetic of anesthetic used should be discussed.

Thank you for your suggestions. We have now discussed EEG/EMG result, limitation of not having physiological measures of anesthesia to control for the anesthesia depth, The possibility that the differences in fos observed may be due to the doses, The possibility that the differences in Fos observed may be due kinetic of anesthetic in the revised manuscript (Lines 308-331, also shown below).

Lines 308-331: "...Our findings indicate that c-Fos expression in the KET group is significantly elevated compared to the ISO group, and the saline group exhibits notably higher c-Fos expression than the home cage group, as seen in Supplementary Figures 2 and 3. Intraperitoneal saline injections in the saline group, despite pre-experiment acclimation with handling and injections for four days, may still evoke pain and stress responses in mice. Subtle yet measurable variations in brain states between the home cage and saline groups were observed, characterized by changes in normalized EEG delta/theta power (home cage: 0.05±0.09; saline: -0.03±0.11) and EMG power (home cage: -0.37±0.34; saline: 0.04±0.13), as shown in Supplementary Figure 1. These changes suggest a relative increase in overall brain activity in the saline group compared to the home cage group, potentially contributing to the higher c-Fos expression. Although the difference in EEG power between the ISO group and the home cage control was not significant, the increase in EEG power observed in the ISO group was similar to that of KET (0.47 ± 0.07 vs 0.59 ± 0.10), suggesting that both agents may induce loss of consciousness in mice. Regarding EMG power, ISO showed a significant decrease in EMG power compared to its control group. In contrast, the KET group showed a lesser reduction in EMG power (ISO: -1.815± 0.10; KET: -0.96 ± 0.21), which may partly explain the higher overall c-Fos expression levels in the KET group. This is consistent with previous studies where ketamine doses up to 150 mg/kg increase delta power while eliciting a wakefulness-like pattern of c-Fos expression across the brain [1]. Furthermore, the observed differences in c-Fos expression may arise in part from the dosages, routes of administration, and their distinct pharmacokinetic profiles. This variation is compounded by the lack of detailed physiological monitoring, such as blood pressure, heart rate, and respiration, affecting our ability to precisely assess anesthesia depth. Future studies incorporating comprehensive physiological monitoring and controlled dosing regimens are essential to further elucidate these relationships and refine our understanding of the effects of anesthetics on brain activity"

1. Lu J, Nelson LE, Franks N, Maze M, Chamberlin NL, Saper CB: Role of endogenous sleep-wake and analgesic systems in anesthesia. J Comp Neurol 2008, 508(4):648-662.

2b. I am confused because Fig 2C seems to show significant decrease in %fos in the hypothalamus, midbrain and cerebellum after KET, while the author responded that " in our analysis, we did not detect regions with significant downregulation when comparing anesthetized mice with controls." Moreover the new figure in the rebuttal in response to reviewer 2 suggests that Ket increases Fos in almost every single region (green vs blue) which is not the conclusion of the paper.

Your concern regarding the apparent discrepancy is well-founded. The inconsistency arose due to an inappropriate data transformation, which affected the interpretation. We have now rectified this by adjusting the data transformation in Figures 2 and 3. Specifically, we have recalculated the log relative c-Fos density values relative to the control group for each brain region. This revision has resolved the issue, confirming that our analysis did not detect any regions with significant downregulation in the anesthetized mice compared to controls. We have also updated the results, discussion, and methods sections of Figures 2 and 3 to accurately reflect these changes and ensure consistency with our findings.

Figure 2.

Whole-brain distributions of c-Fos+ cells induced by ISO and KET. (A) Hierarchical clustering was performed on the log relative c-Fos density data for ISO and KET using the complete linkage method based on the Euclidean distance matrix, with clusters identified by a dendrogram cut-off ratio of 0.5. Numerical labels correspond to distinct clusters within the dendrogram. (B) Silhouette values plotted against the ratio of tree height for ISO and KET, indicating relatively higher Silhouette values at 0.5 (dashed line), which is associated with optimal clustering. (C) The number of clusters identified in each treatment condition at different ratios of the dendrogram tree height, with a cut-off level of 0.5 corresponding to 4 clusters for both ISO and KET (indicated by the dashed line). (D) The bar graph depicts Z scores for clusters in ISO and KET conditions, represented with mean values and standard errors. One-way ANOVA with Tukey's post hoc multiple comparisons. ns: no significance; ***P < 0.001. (E) Z-scored log relative density of c-Fos expression in the clustered brain regions. The order and abbreviations of the brain regions and the numerical labels correspond to those in Figure 2A. The red box denotes the cluster with the highest mean Z score in comparison to other clusters. CTX: cortex; TH: thalamus; HY: hypothalamus; MB: midbrain; HB: hindbrain.

Figure 3.

Similarities and differences in ISO and KET activated c-Fos brain areas. (A) Hierarchical clustering was performed on the log-transformed relative c-Fos density data for ISO and KET using the complete linkage method based on the Euclidean distance matrix, with clusters identified by a dendrogram cut-off ratio of 0.5. (B) Silhouette values are plotted against the ratio of tree height from the hierarchical clustered dendrogram in Figure 3A. (C) The relationship between the number of clusters and the tree height ratio of the dendrogram for ISO and KET, with a cut-off ratio of 0.5 resulting in 3 clusters for ISO and 5 for KET (indicated by the dashed line). (D) The bar graph depicts Z scores for clusters in ISO and KET conditions, represented with mean values and standard errors. One-way ANOVA with Tukey's post hoc multiple comparisons. ns: no significance; ***P < 0.001. (E) Z-scored log relative density of c-Fos expression within the identified brain region clusters. The arrangement, abbreviations of the brain regions, and the numerical labels are in accordance with Figure 3A. The red boxes highlight brain regions that rank within the top 10 percent of Z score values. The white boxes denote brain regions with an Z score less than -2.

1. There are still critical misinterpretations of the PCA analysis. For instance, it is mentioned that " KET is associated with the activation of cortical regions (as evidenced by positive PC1 coefficients in MOB, AON, MO, ACA, and ORB) and the inhibition of subcortical areas (indicated by negative coefficients) " as well as " KET displays cortical activation and subcortical inhibition, whereas ISO shows a contrasting preference, activating the cerebral nucleus (CNU) and the hypothalamus while inhibiting cortical areas. To reduce inter-individual variability." These interpretations are in complete contradiction with the answer 2b above that there was no region that had decreased Fos by either anesthetic.

Thank you for bringing this to our attention. In response to your concerns, we have made significant revisions to our data analysis. We have updated our input data to incorporate log-transformed relative c-Fos density values, normalized against the control group for each brain region, as illustrated in Figures 2 and 3. Instead of PCA, we have applied this updated data to hierarchical clustering analysis. The results of these analyses are consistent with our original observation that neither anesthetic led to a decrease in Fos expression in any region.

1. I still do not understand the rationale for the use of that metric. The use of a % of total Fos makes the data for each region dependent on the data of the other regions which wrongly leads to the conclusion that some regions are inhibited while they are not when looking at the raw data. Moreover, the interdependence of the variable (relative density) may affect the covariance structure which the PCA relies upon. Why not using the PCA on the logarithm of the raw data or on a relative density compared to the control group on a region-per-region basis instead of the whole brain?

Thank you for your insightful suggestion. Following your advice, we have revised our approach and now utilize the logarithm of the relative density compared to the control group on a region-by-region basis. We attempted PCA analyses using the logarithm of the raw data, the logarithm of the Z-score, and the logarithm of the relative density compared to control, but none yielded distinct clusters.

As a result, we employed hierarchical cluster analysis. We then examined the Z-scores of the log-transformed relative c-Fos densities (Figures 2E and 3E) to assess expression levels across clusters. Our analysis revealed that neither ISO nor KET treatments led to a significant suppression of c-Fos expression in the 53 brain regions examined. In the ISO group alone, there were 10 regions that demonstrated relative suppression (Z-score < -2, indicated by white boxes) as shown in Figure 3.

Fig. 2B: it's unclear to me why the regions are connected by a line. Such representation is normally used for time series/within-subject series. What is the rationale for the order of the regions and the use of the line? The line connecting randomly organized regions is meaningless and confusing.

Thank you for your suggestion. We have discontinued the use of PCA calculations and have removed this figure.

Fig 6A. The correlation matrices are difficult to interpret because of the low resolution and arbitrary order of brain regions. I recommend using hierarchical clustering and/or a combination of hierarchical clustering and anatomical organization (e.g. PMID: 31937658). While it is difficult to add the name of the regions on the graph I recommend providing supplementary figures with large high-resolution figures with the name of each brain region so the reader can actually identify the correlation between specific brain regions and the whole brain, Rationale for Metric Choice: Note that I do not dispute the choice of the log which is appropriate, it is the choice of using the relative density that I am questioning.

Thank you for your constructive feedback. In line with your suggestion, we have implemented hierarchical clustering combined with anatomical organization as per the referenced literature. Additionally, we have updated the vector diagrams in Figure 6A to present them with greater clarity.

Furthermore, we have revised our network modular division method based on cited literature recommendations. We used hierarchical clustering with correlation coefficients to segment the network into modules, illustrated in Figure 6—figure supplement 1. Due to the singular module structure of the KET network and the sparsity of intermodular connections in the home cage and saline networks, the assessment of network hub nodes did not employ within-module degree Z-score and participation coefficients, as these measures predominantly underscore the importance of connections within and between modules. Instead, we used degree, betweenness centrality, and eigenvector centrality to detect the hub nodes, as detailed in Figure 6—figure supplement 2. With this new approach, the hub node for the KET condition changed from SS to TeA. Corresponding updates have been made to the results section for Figure 6, as well as to the related discussions and the abstract of our paper.

Figure 6.

Generation of anesthetics-induced networks and identification of hub regions. (A) Heatmaps display the correlations of log c-Fos densities within brain regions (CTX, CNU, TH, HY, MB, and HB) for various states (home cage, ISO, saline, KET). Correlations are color-coded according to Pearson's coefficients. The brain regions within each anatomical category are organized by hierarchical clustering of their correlation coefficients. (B) Network diagrams illustrate significant positive correlations (P < 0.05) between regions, with Pearson’s r exceeding 0.82. Edge thickness indicates correlation magnitude, and node size reflects the number of connections (degree). Node color denotes betweenness centrality, with a spectrum ranging from dark blue (lowest) to dark red (highest). The networks are organized into modules consistent with the clustering depicted in Supplementary Figure 8. Figure 6—figure supplement 1

Figure 6—figure supplement 1.

Hierarchical clustering of brain regions under various conditions: home cage, ISO, saline, and KET. (A) Heatmaps show the relative distances among brain regions assessed in naive mice. Modules were identified by sectioning each dendrogram at a 0.7 threshold. (B) Silhouette scores plotted against the dendrogram tree height ratio for each condition, with optimal cluster definition indicated by a dashed line at a 0.7 ratio. (C) The number of clusters formed at different cutoff levels. At a ratio of 0.7, ISO and saline treatments result in three clusters, whereas home cage and KET conditions yield two clusters. (D) The mean Pearson's correlation coefficient (r) was computed from interregional correlations displayed in Figure 6A. Data were analyzed using one-way ANOVA with Tukey’s post hoc test, ***P < 0.001.

Figure 6—figure supplement 2.

Hub region characterization across different conditions: home cage (A), ISO (B), saline (C), and KET (D) treatments. Brain regions are sorted by degree, betweenness centrality, and eigenvector centrality, with each metric presented in separate bar graphs. Bars to the left of the dashed line indicate the top 20% of regions by rank, highlighting the most central nodes within the network. Red bars signify regions that consistently appear within the top rankings for both degree and betweenness centrality across the metrics.

1. I am still having difficulties understanding Fig. 3.

Panel A: The lack of identification for the dots in panel A makes it impossible to understand which regions are relevant.

Panel B: what is the metric that the up/down arrow summarizes? Fos density? Relative density? PC1/2?

Panel C: it's unclear to me why the regions are connected by a line. Such representation is normally used for time series/within-subject series. What is the rationale for the order of the regions?

Thank you for your patience and for reiterating your concerns regarding Figure 3.

a. In Panel A, we have substituted the original content with a display of hierarchical clustering results, which now clearly marks each brain region. This change aids readers in identifying regions with similar expression patterns and facilitates a more intuitive understanding of the data.

a. Acknowledging that our analysis did not reveal any significantly inhibited brain regions, we have decided to remove the previous version of Panel B from the figure.

b. We have discontinued the use of PCA calculations and have removed this figure to avoid any confusion it may have caused. Our revised analysis focuses on hierarchical clustering, which are presented in the updated figures.

Reviewer #2:

1. Aside from issues with their data transformation (see below), (a) I think they have some interesting Fos counts data in Figures 4B and 5B that indicate shared and distinct activation patterns after KET vs. ISO based anesthesia. These data are far closer to the raw data than PC analyses and need to be described and analyzed in the first figures long before figures with the more abstracted PC analyses. In other words, you need to show the concrete raw data before describing the highly transformed and abstracted PC analyses. (b) This gets to the main point that when selecting brain areas for follow up analyses, these should be chosen based on the concrete Fos counts data, not the highly transformed and abstracted PC analyses.

Thank you for your suggestions.

a. We have added the original c-Fos cell density distribution maps for Figures 2, 3, 4, and 5 in Supplementary Figures 2 and 3 (also shown below). To maintain consistency across the document, we have updated both the y-axis label and the corresponding data in Figures 4B and 5B from 'c-Fos cell count' to 'c-Fos density'.

b. The analyses in Figures 2 and 3 include all brain regions. Figures 4 and 5 present the brain regions with significant differences as shown in Figure 3—figure supplement 1.

Figure 2—figure supplement 1.

The c-Fos density in 53 brain areas for different conditions. (home cage, n = 6; ISO, n = 6 mice; saline, n = 8; KET, n = 6). Each point represents the c-Fos density in a specific brain region, denoted on the y-axis with both abbreviations and full names. Data are shown as mean ± SEM. Brain regions are categorized into 12 brain structures, as indicated on the right side of the graph.

Figure 3—figure supplement 1.

c-Fos density visualization across 201 distinct brain regions under various conditions. The graph depicts the c-Fos density levels for each condition, with data presented as mean and standard error. Brain regions with statistically significant differences are featured in Figures 4 and 5. Brain regions are organized into major anatomical subdivisions, as indicated on the left side of the graph.

1. Now, the choice of data transformation for Fos counts is the most significant problem. First, the authors show in the response letter that not using this transformation (region density/brain density) leads to no clustering. However, they also showed the region-densities without transformation (which we appreciate) and it looks like overall Fos levels in the control group Home (ISO) are a magnitude (~10-fold) higher than those in the control group Saline (KET) across all regions shown. This large difference seems unlikely to be due to a biologically driven effect and seems more likely to be due to a technical issue, such as differences in staining or imaging between experiments. Was the Homecage-ISO experiment or at least the Fos labeling and imaging performed at the same time as for the Saline-Ketamine experiment? Please state the answer to this question in the Results section one way or the other.

a. “Home (ISO) are a magnitude (~10-fold) higher than those in the control group saline (KET) across all regions shown.” We believe you might be indicating that compared to the home cage group (gray), the saline group (blue) shows a 10-fold higher expression (Supplementary Figure 2/3). Indeed, we observed that the total number of c-Fos cells in the home cage group is significantly lower than in the saline group. This difference may be due to reduced sleep during the light-on period (ZT 6- ZT 7.5) in the saline mice or the pain and stress response caused by intraperitoneal injection of saline. We have explained this discrepancy in the discussion section.Line 308-317（also see below）

“…Our findings indicate that c-Fos expression in the KET group is significantly elevated compared to the ISO group, and the saline group exhibits notably higher c-Fos expression than the home cage group, as seen in Supplementary Figures 2 and 3. Intraperitoneal saline injections in the saline group, despite pre-experiment acclimation with handling and injections for four days, may still evoke pain and stress responses in mice. Subtle yet measurable variations in brain states between the home cage and saline groups were observed, characterized by changes in normalized EEG delta/theta power (home cage: 0.05±0.09; saline: -0.03±0.11) and EMG power (home cage: -0.37±0.34; saline: 0.04±0.13), as shown in Figure 1—figure supplement 1. These changes suggest a relative increase in overall brain activity in the saline group compared to the home cage group, potentially contributing to the higher c-Fos expression…”

b. Drug administration and tissue collection for both Homecage-ISO and Saline-Ketamine groups were consistently scheduled at 13:00 and 14:30, respectively. Four mice were administered drugs and had tissues collected each day, with two from the experimental group and two from the control group, to ensure consistent sampling. The 4% PFA fixation time, sucrose dehydration time, primary and secondary antibody concentrations and incubation times, staining, and imaging parameters and equipment (exposure time for VS120 imaging was fixed at 100ms) were all conducted according to a unified protocol.

We have included the following statement in the results section: Line 81-83, “Sample collection for all mice was uniformly conducted at 14:30 (ZT7.5), and the c-Fos labeling and imaging were performed using consistent parameters throughout all experiments. ”

1. Second, they need to deal with this large difference in overall staining or imaging for these two (Home/ISO and Saline/KET) experiments more directly; their current normalization choice does not really account for the large overall differences in mean values and variability in Fos counts (e.g. due to labeling and imaging differences).

3a. I think one option (not perfect but I think better than the current normalization choice) could be z-scoring each treatment to its respective control. They can analyze these z-scored data first, and then in later figures show PC analyses of these data and assess whether the two treatments separate on PC1/2. And if they don't separate, then they don't separate, and you have to go with these results.

3b. Alternatively, they need to figure out the overall intensity distributions from the different runs (if that the main reason of markedly different counts) and adjust their thresholds for Fos-positive cell detection based on this. I would expect that the saline and HC groups should have similar levels of activation, so they could use these as the 'control' group to determine a Fos-positive intensity threshold that gets applied to the corresponding 'treatment' group.

3c. If neither 3a nor 3b is an option then they need to show the outcomes of their analysis when using the untransformed data in the main figures (the untransformed data plots in their responses to reviewer are currently not in the main or supplementary figs) and discuss these as well.

a. Thank you very much for your valuable suggestion. We conducted PCA analysis on the ISO and KET data after Z-scoring them with their respective control groups and did not find any significant separation.

As mentioned in our response to reviewer #1, we have reprocessed the raw data. Firstly, we divided the ISO and KET data by their respective control brain regions and then performed a logarithmic transformation to obtain the log relative c-Fos density. The purpose of this is to eliminate the impact of baseline differences and reduce variability. We then performed hierarchical clustering, and finally, we Z-scored the log relative c-Fos density data. The aim is to facilitate comparison of ISO and KET on the same data dimension (Figure 2 and 3).

b. We appreciate your concerns regarding the detection thresholds for Fos-positive cells. The enclosed images, extracted from supplementary figures for Figures 4 and 5, demonstrate notable differences in c-Fos expression between saline and home cage groups in specific brain regions. These regions exhibit a discernible difference in staining intensity, with the saline group showing enhanced c-Fos expression in the PVH and PVT regions compared to the home cage group. An examination of supplementary figures for Figures 4 and 5 shows that c-Fos expression in the home cage group is consistently lower than in the saline group. This comparative analysis confirms that the discrepancies in c-Fos levels are not due to varying detection thresholds.

b. We have added the corresponding original data graphs to Supplementary Figures 2 and 3, and discussed the potential reasons for the significant differences between these groups in the discussion section (also shown below).

Lines 308-317: "...Our findings indicate that c-Fos expression in the KET group is significantly elevated compared to the ISO group, and the saline group exhibits notably higher c-Fos expression than the home cage group, as seen in Supplementary Figures 2 and 3. Intraperitoneal saline injections in the saline group, despite pre-experiment acclimation with handling and injections for four days, may still evoke pain and stress responses in mice. Subtle yet measurable variations in brain states between the home cage and saline groups were observed, characterized by changes in normalized EEG delta/theta power (home cage: 0.05±0.09; saline: -0.03±0.11) and EMG power (home cage: -0.37±0.34; saline: 0.04±0.13), as shown in Figure 3—figure supplement 1. These changes suggest a relative increase in overall brain activity in the saline group compared to the home cage group, potentially contributing to the higher c-Fos expression.…”

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Ketamine (KET) and isoflurane (ISO) are two widely used general anesthetics, yet their distinct and shared neurophysiological mechanisms remain elusive. In this study, we conducted a comparative analysis of KET and ISO effects on c-Fos expression across the brain, utilizing principal component analysis (PCA) and c-Fos-based functional network analysis to evaluate the responses of individual brain regions to each anesthetic. Our findings demonstrate that KET significantly activates cortical and subcortical arousal-promoting nuclei, with the somatosensory cortex (SS) serving as a hub node, corroborating the top-down general anesthesia theory for dissociative anesthesia. In contrast, ISO activates the nuclei in the hypothalamus and brainstem, with the locus coeruleus (LC) as a hub node, implying a bottom-up mechanism for anesthetic-induced unconsciousness. Notably, the coactivation of sleep-wakefulness regulation, analgesia-related, neuroendocrine-related nuclei (e.g., prelimbic area (PL) and infralimbic areas (ILA), and the anterior paraventricular nucleus (aPVT), Edinger-Westphal nucleus (EW), locus coeruleus (LC), parabrachial nucleus (PB), solitary tract nucleus (NTS)) by both anesthetics underscores shared features such as unconsciousness, analgesia, and autonomic regulation, irrespective of their specific molecular targets. In conclusion, our results emphasize the distinct actions of KET and ISO while also uncovering the commonly activated brain regions, thus contributing to the advancement of our understanding of the mechanisms underlying general anesthesia.

Revised version: Ketamine (KET) and isoflurane (ISO) are two widely used general anesthetics, yet their distinct and shared neurophysiological mechanisms remain elusive. In this study, we conducted a comparative analysis of KET and ISO effects on c-Fos expression across the brain, utilizing hierarchical clustering and c-Fos-based functional network analysis to evaluate the responses of individual brain regions to each anesthetic. Our findings demonstrate that KET significantly activates cortical and subcortical arousal-promoting nuclei, with the temporal association areas (TEa) serving as a hub node, corroborating the top-down general anesthesia theory for dissociative anesthesia. In contrast, ISO activates the nuclei in the hypothalamus and brainstem, with the locus coeruleus (LC) as a hub node, implying a bottom-up mechanism for anesthetic-induced unconsciousness. Notably, the coactivation of arousal-related nuclei, analgesia-related, neuroendocrine-related nuclei (e.g., prelimbic area (PL) and infralimbic areas (ILA), and the anterior paraventricular nucleus (aPVT), Edinger-Westphal nucleus (EW), locus coeruleus (LC), parabrachial nucleus (PB), solitary tract nucleus (NTS) by both anesthetics underscores shared features such as unconsciousness, analgesia, and autonomic regulation, irrespective of their specific molecular targets. In conclusion, our results emphasize the distinct actions of KET and ISO while also uncovering the commonly activated brain regions, thus contributing to the advancement of our understanding of the mechanisms underlying general anesthesia.

Introduction

Earlier version: Despite considerable investigation into the molecular targets, neural circuits, and functional connectivity associated with various anesthetics, our comprehension of their effects on overall brain activity continues to be limited and incomplete [1]. At the molecular level, ketamine (KET) and isoflurane (ISO) interact with N-methyl-D-aspartate (NMDA) and gamma-aminobutyric acid type A (GABAa) receptors, respectively, modulating neuronal excitability and ultimately leading to a loss of consciousness [2]. In systems neuroscience, the neural mechanisms of anesthetic induced unconsciousness involve both top-down and bottom-up processes [3, 4]. As evidenced by in vivo electrophysiology or functional magnetic resonance imaging (fMRI) studies, the top-down paradigm illustrates that anesthetics induce unconsciousness by disrupting corticocortical and corticothalamic circuits responsible for neural information integration, while peripheral sensory information can still be conveyed to the primary sensory cortex [5, 6]. The bottom-up approach, exemplified by ISO, involves the activation of sleep-promoting nuclei like ventral lateral preoptic nucleus (VLPO) and inhibition of arousal centers in the brainstem and diencephalon, supporting the shared circuits of sleep and anesthesia [7, 8]. However, the limited spatial resolution of fMRI studies and the inability of EEG to capture specific brainstem nuclei hinder the acquisition of comprehensive whole-brain information. Although a substantial body of knowledge has been amassed, our understanding of the reciprocal responses among different brain regions during general anesthesia remains relatively sparse and fragmented. To bridge these gaps, further investigation using advanced techniques that can capture the whole-brain dynamics is needed to elucidate the complex interactions and shared mechanisms between various anesthetics.

Revised version: Despite considerable investigation into the molecular targets, neural circuits, and functional connectivity associated with various anesthetics, our comprehension of their effects on overall brain activity continues to be limited and incomplete (Hemmings et al., 2019). At the molecular level, ketamine (KET) and isoflurane (ISO) interact with N-methyl-D-aspartate (NMDA) and gamma-aminobutyric acid type A (GABAa) receptors, respectively, modulating neuronal excitability and ultimately leading to a loss of consciousness (Franks, 2008). In systems neuroscience, the neural mechanisms of anesthetic induced unconsciousness involve both top-down and bottom-up processes (Mashour, 2014; Mashour and Hudetz, 2017). As evidenced by in vivo electrophysiology or functional magnetic resonance imaging (fMRI) studies, the top-down paradigm illustrates that anesthetics induce unconsciousness by disrupting corticocortical and corticothalamic circuits responsible for neural information integration, while peripheral sensory information can still be conveyed to the primary sensory cortex (Schroeder et al., 2016; Lee et al., 2013). The bottom-up approach, exemplified by ISO, involves the activation of sleep-promoting nuclei like ventral lateral preoptic nucleus (VLPO) and inhibition of arousal centers in the brainstem and diencephalon, supporting the shared circuits of sleep and anesthesia (Moore et al., 2012; Nelson et al., 2002). However, the limited spatial resolution of fMRI studies and the inability of EEG to capture specific brainstem nuclei hinder the acquisition of comprehensive whole-brain information. Although a substantial body of knowledge has been amassed, our understanding of the reciprocal responses among different brain regions during general anesthesia remains relatively sparse and fragmented. To bridge these gaps, further investigation using advanced techniques that can capture the whole-brain dynamics is needed to elucidate the complex interactions and shared mechanisms between various anesthetics.

Earlier version: Neuronal extracellular stimulation typically results in the elevation of adenosine 3’,5’-cyclic monophosphate (cAMP) levels and calcium influx, ultimately leading to the upregulation of immediate early genes (IEGs) such as c-fos [9, 10]. The translation product of c-fos, c-Fos protein, offers single-cell spatial resolution and has been utilized as a biomarker to identify anesthetic-activated brain regions [11]. Previous investigations of c-Fos expression throughout the brain demonstrated that GABAergic agents inhibited cortical activity while concurrently activating subcortical brain regions, including the VLPO, median preoptic nucleus (MnPO), lateral septal nucleus (LS), Edinger-Westphal nucleus (EW), and locus coeruleus (LC) [12–15]. In contrast, KET was shown to provoke wake-like c-Fos expression and intense augmentation of c-Fos expression in various brain regions at clinical dosages (75-100 mg/kg) [13]. It is important to note that these experiments administered KET at lights-on and GABAa receptor agonists at lights-off, potentially introducing circadian influences for direct comparison of ISO and KET. Moreover, it has been revealed that state of general anesthesia is not determined by activity in individual brain areas, but emerges as a global change within the brain. This change involves the activation of lateral habenular nucleus (LHb), VLPO, supraoptic nucleus (SON), and central amygdaloid nucleus (CeA), which are essential for anesthetics induced sedation, unconsciousness, or analgesia [7, 16–18]. However, brain-wide mapping and comparison of distinct anesthetic (KET and ISO) activated nuclei at a cellular level have not been fully elucidated.

Revised version: Neuronal extracellular stimulation typically results in the elevation of adenosine 3’,5’-cyclic monophosphate (cAMP) levels and calcium influx, ultimately leading to the upregulation of immediate early genes (IEGs) such as c-fos (Yap and Greenberg, 2018; Morgan and Curran, 1989). The translation product of c-fos, c-Fos protein, offers single-cell spatial resolution and has been utilized as a biomarker to identify anesthetic-activated brain regions (Zhang et al., 2022). Previous investigations of c-Fos expression throughout the brain demonstrated that GABAergic agents inhibited cortical activity while concurrently activating subcortical brain regions, including the VLPO, median preoptic nucleus (MnPO), lateral septal nucleus (LS), Edinger-Westphal nucleus (EW), and locus coeruleus (LC) (Smith et al., 2016; Lu et al., 2008; Yatziv et al., 2020; Han et al., 2014). In contrast, KET was shown to provoke wake-like c-Fos expression and intense augmentation of c-Fos expression in various brain regions at clinical dosages (75-100 mg/kg) (Lu et al., 2008). It is important to note that these experiments administered KET at lights-on and GABAa receptor agonists at lights-off, potentially introducing circadian influences for direct comparison of ISO and KET. Moreover, it has been revealed that state of general anesthesia is not determined by activity in individual brain areas, but emerges as a global change within the brain. This change involves the activation of lateral habenular nucleus (LHb), VLPO, supraoptic nucleus (SO), and central amygdaloid nucleus (CeA), which are essential for anesthetics induced sedation, unconsciousness, or analgesia (Moore et al., 2012; Gelegen et al., 2018; Hua et al., 2020; Jiang-Xie et al., 2019). However, brain-wide mapping and comparison of distinct anesthetic (KET and ISO) activated nuclei at a cellular level have not been fully elucidated.

Earlier version: In this study, we examined the distribution of nuclei activated by ISO and KET in 987 brain regions using immunochemical labeling and a customized MATLAB software package, which facilitated signal detection and registration to the Allen Mouse Brain Atlas reference [19]. We compared whole-brain c-Fos expression induced by KET and ISO through principal component analysis (PCA) and calculated inter-regional correlations by determining the covariance across subjects for each brain region pair. We then extracted significantly positively correlated brain regions to construct functional networks and performed graph theory-based network analyses to identify hub nodes. Our results uncovered distinct yet overlapping whole-brain activation patterns for KET and ISO.

Revised version: In this study, we examined the distribution of nuclei activated by ISO and KET in 984 brain regions using immunochemical labeling and a customized MATLAB software package, which facilitated signal detection and registration to the Allen Mouse Brain Atlas reference (Ma et al., 2021). We compared whole-brain c-Fos expression induced by KET and ISO through hierarchical clustering and calculated inter-regional correlations by determining the covariance across subjects for each brain region pair. Significantly positively correlated regions were then extracted to construct functional networks and graph theory-based analyses were performed to identify hub nodes. Our results revealed distinct yet overlapping whole-brain activation patterns for KET and ISO.

Results

A comparison of the activation patterns of c-Fos in 53 brain areas in response to ISO and KET

Added in the revised version: To examine the pattern of c-Fos expression throughout the brain, 1.5% ISO was continuously ventilated, or 100 mg/kg KET was administered 90 minutes before harvesting (Figure 1A). Sample collection for all mice was uniformly conducted at 14:30 (ZT7.5), and the c-Fos labeling and imaging were performed using consistent parameters throughout all experiments. To quantitatively assess brain states across the four groups, we utilized electroencephalography (EEG) and electromyography (EMG) recordings (Figure 1—figure supplement 1). The depth of anesthesia was evaluated based on the power ratio of the EEG within the delta (0.5–4 Hz) and theta (6–10 Hz) frequency bands, combined with the changes in the EMG (Luo et al., 2023). Our findings showed that ISO and KET elevated EEG power ratio while reducing EMG power, suggesting both drugs induced a loss of consciousness. The raw images of brain slices were aligned to the Allen Mouse Brain Atlas (Figure 1B). Based on the standard mouse atlas available at http://atlas.brain-map.org/, the mouse brain was segmented into nine hierarchical levels, totaling 984 regions. The primary level consists of grey matter, the secondary of the cerebrum, brainstem, and cerebellum, and the tertiary includes regions like the cerebral cortex and cerebellar nuclei, among others, with some regions extending to the 8th and 9th levels. The fifth level comprises 53 subregions, with detailed expression levels and their respective abbreviations presented in Figure 2—figure supplement (Do et al., 2016).

Removed in the revised version: To examine the pattern of c-Fos expression throughout the brain, 1.5% ISO was continuously ventilated, or 100 mg/kg KET was administered 90 minutes before harvesting (Figure 1A). Concurrently, to quantitatively assess brain states across the four groups, we utilized electroencephalography (EEG) and electromyography (EMG) recordings (Figure 1—figure supplement 1). The depth of anesthesia was evaluated based on the power ratio of the EEG within the delta (0.5–4 Hz) and theta (6–10 Hz) frequency bands, combined with the changes in the EMG [20]. Our findings showed that ISO and KET notably elevated EEG power ratio while reducing EMG power, suggesting both drugs induced a loss of consciousness. The raw images of brain slices were aligned to the Allen Mouse Brain Atlas (Figure 1B). Based on the standard mouse atlas available at http://atlas.brain-map.org/, the mouse brain was segmented into nine hierarchical levels, totaling 987 regions. The primary level consists of grey matter, the secondary of the cerebrum, brainstem, and cerebellum, and the tertiary includes regions like the cerebral cortex and cerebellar nuclei, among others. Some regions extend to the 8th and 9th levels. Our initial statistical comparison was conducted on the fifth level, which is composed of 53 subregions, as detailed in Supplementary Table 1 (Figure 1B) [21]. To differentiate c-Fos expression patterns across groups, we calculated the c-Fos density for each brain region, divided it by the overall brain c-Fos density, and then applied a logarithmic transformation. The transformed values were used for principal component analysis (PCA) (Figure 2A). The KET and ISO groups were distinctly delineated by the top two principal components. Notably, the KET-induced c-Fos expression manifested a positive association with principal component 1 (PC1), clustering around +2.5, which represented 30.72% of the total variance. In contrast, it exhibited a negative correlation with principal component 2 (PC2), clustering near −2.0, which contributed to 25.89% of the variance. On the other hand, ISO displayed a positive correlation with PC2, clustering at approximately +2.5, while its association with PC1 was neutral, near 0. Figure 2B illustrates the component coefficients of PC1 and PC2 for each brain region. Dots represent the top 25% of brain regions based on the absolute values of their component coefficients, signifying the most influential regions. KET, predominantly positively correlated with PC1, is associated with the activation of cortical regions (as evidenced by positive PC1 coefficients in MOB, AON, MO, ACA, and ORB) and the inhibition of subcortical areas (indicated by negative coefficients). On the other hand, ISO, with its notable association with PC2, showed top 25% positive coefficients in PALc, LSX of CNU, cortical ILA, and hypothalamic PVZ, while its negative coefficients were found primarily in cortical regions. These observations underscore that KET and ISO exhibit distinct patterns of brain region activation. KET displays cortical activation and subcortical inhibition, whereas ISO shows a contrasting preference, activating the cerebral nucleus (CNU) and the hypothalamus while inhibiting cortical areas. To reduce inter-individual variability, we normalized c-Fos+ cell counts in each brain area by dividing them by the total c-Fos+ cells in the entire brain (Figure 2C), Subsequently, we directly compared the differences between ISO and KET with their respective control groups. In line with PCA results, KET significantly activated the isocortex (16.55 ± 2.31 vs. 40.13 ± 1.97, P < 0.001) and reduced the proportion of c-Fos+ cells in the hypothalamus (14.88 ± 2.18 vs. 4.03 ± 0.57, P = 0.001) and midbrain (21.87 ± 2.74 vs. 6.48 ± 1.06, P = 0.005) compared to the saline group. In contrast, the proportion of c-Fos+ cells induced by ISO were mainly observed in the striatum and hypothalamus compared to the home cage group, specifically in the periventricular zone (PVZ) (2.330 ± 0.759 vs. 6.63 ± 0.84, P = 0.015) and lateral zone (LZ) (2.82 ± 0.3 vs. 5.59 ± 0.75, P = 0.002) of the hypothalamus, as well as STRv (0.59 ± 0.10 vs. 2.77 ± 0.40, P < 0.002) and LSX (1.93 ± 0.59 vs. 5.11 ± 0.41, P = 0.004) of the striatum (Figure 2C and Table 1). In summary, our findings reveal that KET predominantly affects cortical regions, whereas isoflurane mainly targets subcortical areas, specifically the striatum, and hypothalamus. These results demonstrate distinct patterns of brain region activation for each anesthetic agent.

Brain-wide quantification of c-Fos expression.

Earlier version: (A) Schematic representation of the habituation protocol typically used to acclimate mice. After being exposed to anesthetics for 90 minutes, the mice were euthanized. (B) Steps for data processing. Example of brain section registration to a corresponding coronal section from the Allen Brain Atlas. For Atlas rotation, the Allen reference atlas was rotated to mimic the slice angle of the experimental brain. Image registration maps the original coronal image (upper panel) to the corresponding Allen mouse brain atlas (lower panel). The registration module applies several geometric transformations (translation, rotation, and scaling) to optimize the matching of the original image to the anatomical structures. Fluorescence signals were detected from the original image (upper panel), and once detected, they were projected onto the Allen Mouse Brain Atlas for quantification and network analysis by means of the detected signals labeled with yellow boxes.

Revised version: (A)Schematic representation of the habituation protocol typically used to acclimate mice. After being exposed to anesthetics for 90 minutes, the mice were euthanized. (B) Steps for data processing. Example of brain section registration to a corresponding coronal section from the Allen Brain Atlas. For Atlas rotation, the Allen reference atlas was rotated to mimic the slice angle of the experimental brain. Image registration maps the original coronal image (upper panel) to the corresponding Allen mouse brain atlas (lower panel). The registration module applies several geometric transformations (translation, rotation, and scaling) to optimize the matching of the original image to the anatomical structures. Fluorescence signals detected in the original image were projected onto the Allen Mouse Brain Atlas for quantification. Finally, the processed data undergo hierarchical clustering and network analysis to investigate the patterns of c-Fos expression and central network nodes. Figure 1—figure supplement 1. EEG and EMG power change after each treatment. The box represents the 25th-75th percentiles; the central line denotes the median; whiskers indicate maximum and minimum values. n = 6, 6, 8, 6 for the home cage, ISO, saline, and KET, respectively. (A) Normalized change in EEG power: ISO vs KET, P > 0.99; Saline vs KET, P = 0.01; Home cage vs ISO, P = 0.11. (B) Normalized change in EMG power: ISO vs KET, P = 0.36; Saline vs KET, P = 0.30; Home cage vs ISO, P = 0.02. Analyses were conducted using the Kruskal-Wallis test, followed by Dunn’s multiple comparisons tests.

Earlier version: Whole-brain distributions of c-Fos+ cells induced by ISO, KET, and control conditions.

Revised version: Whole-brain distributions of c-Fos+ cells induced by ISO and KET.

Added in the revised version: (A) Hierarchical clustering was performed on the log relative c-Fos density data for ISO and KET using the complete linkage method based on the Euclidean distance matrix, with clusters identified by a dendrogram cut-off ratio of 0.5. Numerical labels correspond to distinct clusters within the dendrogram. (B) Silhouette values plotted against the ratio of tree height for ISO and KET, indicating relatively higher Silhouette values at 0.5 (dashed line), which is associated with optimal clustering. (C) The number of clusters identified in each treatment condition at different ratios of the dendrogram tree height, with a cut-off level of 0.5 corresponding to 4 clusters for both ISO and KET (indicated by the dashed line). (D) The bar graph depicts Z scores for clusters in ISO and KET conditions, represented with mean values and standard errors. One-way ANOVA with Tukey’s post hoc multiple comparisons. ns: no significance; ***P < 0.001. (E) Z-scored log relative density of c-Fos expression in the clustered brain regions. The order and abbreviations of the brain regions and the numerical labels correspond to those in Figure 2A. The red box denotes the cluster with the highest mean Z score in comparison to other clusters. CTX: cortex; TH: thalamus; HY: hypothalamus; MB: midbrain; HB: hindbrain. Figure 2—figure supplement 1. The c-Fos density in 53 brain areas for different conditions. (home cage, n = 6; ISO, n = 6 mice; saline, n = 8; KET, n = 6). Each point represents the c-Fos density in a specific brain region, denoted on the y-axis with both abbreviations and full names. Data are shown as mean ± SEM. Brain regions are categorized into 12 brain structures, as indicated on the right side of the graph.

Added in the revised version: To determine the effects of ISO and KET on neural activity, we normalized c-Fos densities across brain regions to their control group averages region by region and log-transformed the data. Hierarchical clustering of these values revealed distinct patterns of drug-induced changes (Figure 2A), with both ISO and KET showing optimal clustering at a 0.5 cut-off value, as confirmed by high silhouette coefficients (Figure 2B). Concurrently, we computed the number of clusters formed at each potential cut-off value and found that a 0.5 threshold resulted in four distinct clusters for both ISO and KET treatments (Figure 2C). After hierarchical clustering, we maintained the sequence of brain regions and substituted log relative c-Fos densities with Z-scores for both ISO and KET (Figure 2E). While no significantly inhibited brain regions were detected in the ISO group (as no Z scores fell below −2), clusters 3 and 4 showed expression levels similar to or lower than those in the control group. These clusters are involved in sensory integration, motor coordination, and higher-order cognitive processing, with cluster 3 having a mean Z score of 0.27 (n=17) and cluster 4 a mean Z score of −1.1 (n=2), as shown in Figures 2D and 2E. The most pronounced upregulation of ISO induced c-Fos expression was observed in the second cluster (n=13, mean Z score=4.7), which includes brain regions responsible for a variety of functions: the periventricular zone (PVZ) and lateral zone (LZ) of the hypothalamus, vital for endocrine and autonomic system regulation; the ventral striatum (STRv), pallidum (PAL), and lateral septal complex (LSX) of the striatum, key in reward and motivation; the tenia tecta (TT) and piriform area (PIR), essential for olfactory processing; and the infralimbic area (ILA) of the cortex, significant for emotional and cognitive functions. This collective upregulation in cluster 2 underlines ISO’s impact on hypothalamic regulation, reward mechanisms, and olfactory processing. KET administration elicited a significant activation in a broad array of brain regions, with Cluster 4 displaying the most substantial upregulation in expression. This cluster incorporates cortical areas central to sensory processing and emotional regulation, such as the visceral area (VISC), anterior cingulate area (ACA), and primary somatosensory cortex (SS) — as well as cerebral nuclei integral for motor coordination and motivational behaviors — specifically, the dorsolateral and ventrolateral pallidum (PALd, PALv) and striatum (STRd, STRv) (Figures 2D and 2E). ISO predominantly modulates brain regions associated with hormonal regulation, reward processing, and olfaction, with marked activation in the hypothalamic and striatal regions. In contrast, KET induces widespread activation across the brain, with heightened activity in a key cluster that includes cortical and cerebral nuclei governing complex sensory, emotional, and motor processes, thus highlighting its notable influence on higher-order functions.

Removed in the revised version: (A) Scatter plot of four conditions in the space spanned by the first two principal components (PC1 versus PC2) of c-Fos density from 53 brain regions. (B) Line plot of coefficients of PC1 and PC2. Dots represent regions with absolute values larger than 75 percentiles. (C) The normalized c-Fos+ cells in 53 brain areas (Home cage, n = 6; ISO, n= 6 mice; Saline, n = 8; KET, n = 6). Brain areas are grouped into 12 generalized, color-coded brain structures. Abbreviations of the 53 brain areas and percentages of c-Fos+ cells are listed in supplementary table 1. Error bar, mean ± SEM. Mann-Whitney tests were corrected using the false discovery rate approach with a Q-value set at 0.05. Home cage vs ISO: *P<0.05, **P<0.01, ***P<0.001. Saline vs KET: #, P<0.05, ##, P<0.01, ###, P<0.001.

Similarities and differences in ISO and KET activated c-Fos brain areas

Added in the revised version: To enhance our analysis of c-Fos expression patterns induced by ISO and KET, we expanded our study to 201 subregions. After log-transforming the c-Fos densities relative to controls, we applied hierarchical clustering to uncover patterns of brain activity (Figure 3A), with detailed expression levels in Figure 3—figure supplement 1. ISO demonstrates a more concentrated influence on brain activity with consistently higher clustering quality and fewer clusters than KET, suggesting a focused impact on specific brain functions (Figures 3B and 3C). Figure 3E maintains the order from hierarchical clustering and marks both the most activated brain regions in the top 10% of Z-scores with red boxes and areas of potential suppression with Z-scores below −2 with white boxes. ISO hierarchical clustering identified three distinct clusters, with cluster 3 exhibiting the highest mean Z-score of 7.27. This cluster encompasses regions like the supraoptic (SO), ventrolateral preoptic (VLPO), tuberal (TU), and central amygdala (CEA). Notably, the SO, VLPO, and CEA have been previously reported to play critical roles in the mechanisms of anesthesia (Moore et al., 2012; Hua et al., 2020; Jiang-Xie et al., 2019). In contrast, cluster 1 showed a relative decrease in expression (n= 91, mean Z-score: −0.54), indicating that nearly half of the brain areas had reduced c-Fos expression compared to the control group. Cluster 2 was overall moderately activated, with significant activity in the top 10% of regions including the anterior hypothalamic (ARH) and paraventricular (PVH) nuclei—key for hypothalamic functions—and cortical areas such as the infralimbic (ILA), dorsal peduncular (DP), and temporal (TT) cortices, which are involved in emotional and cognitive processing. Striatal and pallidal regions, particularly the olfactory tubercle (OT) and nucleus accumbens (ACB), as well as midbrain and hindbrain structures like the area postrema (AP) and nucleus of the solitary tract (NTS), showed the most pronounced activation. In the hierarchical clustering analysis of KET exposure, Cluster 2 with a mean Z-score of −0.01 was similar to the control group in terms of activation, while the other four clusters showed increased activation. Cluster 5, showing the most significant activation, is composed mainly of cortical regions critical for sensory processing, spatial memory, navigation, emotional regulation, and cognitive functions. It also includes cerebral nuclei involved in motor control and reward pathways, alongside essential thalamic and hindbrain areas for memory and sensory transmission. This composition, with a predominant focus on cortical regions, suggests that KET chiefly modulates higher-order brain functions, while also affecting basic neurological processes through its influence on cerebral nuclei, the thalamus, and the hindbrain. In conclusion, while ISO’s activation of specific areas like SO, VLPO, TU, and CeA in cluster 3 indicates a targeted influence, KET’s extensive engagement with cortical areas in cluster 5, along with its effects on striatal and thalamic regions, underscores its profound and widespread modulation of both higher-order and fundamental neural functions.

Removed in the revised version: To elucidate the c-Fos expression patterns induced by ISO and KET more clearly, we refined our focus from the initial 53 brain regions to a more granular examination of subregions, culminating in a comprehensive assessment of 201 distinct areas. We conducted a PCA on the logarithmic ratio of c-Fos density across these regions to discern variance among the groups (Figure 3). The KET group predominantly clustered around +7 on PC1 and approximately 0 on PC2, whereas the ISO group clustered around +4 on PC2 and −3 on PC1. KET displayed a strong positive correlation with PC1, accounting for 28.97% of the variance. In contrast, ISO negatively correlated with PC1 and positively with PC2, which accounted for 15.54% of the variance (Figure 3A). These findings underscore the distinct features between the ISO and KET groups, with the top two principal components effectively separating them. Further analysis of PC coefficients in each brain region showed (Figure 3B) that KET, positively associated with PC1, mainly activates the cerebral cortex (CTX) areas including the visceral area (VISC), claustrum (CLA), dorsal peduncular area (DP), orbital area (ORB), and temporal association areas (TEA), while inhibiting the hypothalamus, midbrain, and hindbrain. In contrast, ISO is negatively correlated with PC1, suggesting that the cortical areas activated by KET are inhibited by ISO. Additionally, ISO has a positive association with PC2 and is predominantly related to the activation of hypothalamic regions, such as the ventrolateral preoptic nucleus (VLPO), suprachiasmatic nucleus (SCH), tuberal nucleus (TU), medial preoptic area (MPO), and supraoptic nucleus (SON), as well as certain nuclei in the central nucleus (CNU). KET primarily activates regions within the cerebral cortex (CTX) and inhibits areas within the hypothalamus, midbrain, and hindbrain, aligning with the top-down mechanism suggesting anesthetic agents influence cortical and thalamocortical circuits [4]. Conversely, ISO activates hypothalamic areas while suppressing the cortex, midbrain, and hindbrain, consistent with the bottom-up mechanism. Figure 3C summarizes these distinctive responses, implying that anesthetic agents like KET and ISO modulate consciousness by targeting particular brain regions and neural circuits.

Similarities and differences in ISO and KET activated c-Fos brain areas.

Added in the revised version: (A) Hierarchical clustering was performed on the log-transformed relative c-Fos density data for ISO and KET using the complete linkage method based on the Euclidean distance matrix, with clusters identified by a dendrogram cut-off ratio of 0.5. Brain region labels are provided in Supplementary Figures 4 and 5. (B) Silhouette values are plotted against the ratio of tree height from the hierarchical clustered dendrogram in Figure 3A. (C) The relationship between the number of clusters and the tree height ratio of the dendrogram for ISO and KET, with a cut-off ratio of 0.5 resulting in 3 clusters for ISO and 5 for KET (indicated by the dashed line). (D) The bar graph depicts Z scores for clusters in ISO and KET conditions, represented with mean values and standard errors. One-way ANOVA with Tukey’s post hoc multiple comparisons. ns: no significance; ***P < 0.001. (E) Z-scored log relative density of c-Fos expression within the identified brain region clusters. The arrangement, abbreviations of the brain regions, and the numerical labels are in accordance with Figure 3A. The red boxes highlight brain regions that rank within the top 10 percent of Z score values. The white boxes denote brain regions with an Z score less than −2. Figure 3—figure supplement 1. c-Fos density visualization across 201 distinct brain regions under various conditions. The graph depicts the c-Fos density levels for each condition, with data presented as mean and standard error. Brain regions with statistically significant differences are featured in Figures 4 and 5. Brain regions are organized into major anatomical subdivisions, as indicated on the left side of the graph. Figure 3—figure supplement 2. Region labels for the hierarchical clustering of the ISO group in Figure 3A.

Removed in the revised version: (A) Scatter plot of four conditions in the space spanned by the first two principal components (PC1 versus PC2) of normalized c-Fos density from 201 brain regions. (B) Line plot of coefficients of PC1 and PC2. Dots represent regions with absolute values larger than 75 percentiles. (C) Activation patterns of ISO and KET in different parts of brain regions were summarized. ISO was predominantly characterized by increased activation of hypothalamic and cerebral nuclei while reducing activity in the cortex and midbrain, indicating that ISO primarily exerts its effects through a bottom-up mechanism. On the other hand, KET primarily activated the cortex and inhibited hypothalamus and midbrain activity, suggesting that KET regulates consciousness through a top-down mechanism.

Identification of brain regions activated by KET

Earlier version: We further employed quantitative methods to identify the brain regions activated by KET across the entire brain. Our findings concur with previous studies: numerous cortical regions associated with somatosensory, auditory, visual, and movement were activated (Figure 4A and C, Supplementary Figure 1, Supplementary Table 2). Additionally, we identified several innovative observations that enrich the current understanding in this field. To provide a clearer overview of our findings, we categorized the activated brain areas into different functional groups: (1) Arousal and REM sleep regulation: Several nuclei associated with arousal were activated, including the prelimbic area (PL), infralimbic cortex (ILA), paraventricular nucleus of the thalamus (PVT), dorsal raphe (DR), and sublaterodorsal nucleus (SLD). Prior evidence indicates that the PL/ILA and PVT regions play a role in regulating arousal in both cortical and subcortical areas[22, 23]. Additionally, dopaminergic neurons within the DR have been identified as vital components of wake-promoting pathways [24]. In the context of REM sleep, the sublaterodorsal nucleus (SLD) stands out for its role in stabilization[25]. (2) Pain modulation, KET significantly activated pain-related areas such as the anterior cingulate cortex (ACA) in the cortex [26], the anterior pretectal nucleus (APN) in the thalamus, which is known to be involved in managing chronic pain [27], and the anterior periaqueductal gray (PAG) region and medial prefrontal cortex (mPFC), both part of the endogenous pain inhibitory pathway [28, 29]. Additionally, the activation of the locus coeruleus (LC) in the midbrain may contribute to KET’s analgesic effects [30]. (3) Neuroendocrine regulation: The paraventricular hypothalamic nucleus (PVH) and supraoptic nucleus (SON), which are neuroendocrine-related regions, were also activated [18, 31]. (4) Movement: Subcortical nuclei associated with movement, such as the subthalamic nucleus (STN) and nucleus incertus (NI), were prominently activated by KET administration [32, 33]. (5) Connectivity: We observed significant activation of the nucleus reuniens (Re) located in the thalamus, which receives substantial afferent input from limbic structures and serves as a connector linking the hippocampus to the medial prefrontal cortex [34]. While our functional categorization offers a simplified overview, it’s worth noting that each activated brain region may have multifaceted roles beyond this classification. In summary, our study identified extensive activation of cortical and subcortical nuclei during KET anesthesia, encompassing regions related to arousal, pain modulation, neuroendocrine regulation, movement, and connectivity.

Revised version: We further employed quantitative methods to identify the brain regions activated by KET across the entire brain. Our findings concur with previous studies: numerous cortical regions associated with somatosensory, auditory, visual, and movement were activated (Figure 4A and C, Figure 4—figure supplement 1, Figure 4—source data 1). Additionally, we identified several innovative observations that enrich the current understanding in this field. To provide a clearer overview of our findings, we categorized the activated brain areas into different functional groups: (1) Arousal and REM sleep regulation: Several nuclei associated with arousal were activated, including the prelimbic area (PL), infralimbic cortex (ILA), paraventricular nucleus of the thalamus (PVT), dorsal raphe (DR), and sublaterodorsal nucleus (SLD). Prior evidence indicates that the PL/ILA and PVT regions play a role in regulating arousal in both cortical and subcortical areas (Mashour et al., 2022; Gao et al., 2020). Additionally, dopaminergic neurons within the DR have been identified as vital components of wake-promoting pathways (Cho et al., 2017). In the context of REM sleep, the sublaterodorsal nucleus (SLD) stands out for its role in stabilization (Feng et al., 2020). (2) Pain modulation: KET significantly activated pain-related areas such as the anterior cingulate cortex (ACA) in the cortex (Bliss et al., 2016), the anterior pretectal nucleus (APN) in the thalamus, which is known to be involved in managing chronic pain (Villarreal et al., 2003), and the anterior periaqueductal gray (PAG) region and medial prefrontal cortex (mPFC), both part of the endogenous pain inhibitory pathway (Cheriyan and Sheets, 2018; Yeung et al., 1977). Additionally, the activation of the locus coeruleus (LC) in the midbrain may contribute to KET’s analgesic effects (Llorca-Torralba et al., 2021). (3) Neuroendocrine regulation: The paraventricular hypothalamic nucleus (PVH) and supraoptic nucleus (SO), which are neuroendocrine-related regions, were also activated (Jiang-Xie et al., 2019; Duncan et al., 1998). (4) Movement: Subcortical nuclei associated with movement, such as the subthalamic nucleus (STN) and nucleus incertus (NI), were prominently activated by KET administration (Ma et al., 2017; Musacchio et al., 2017). (5) Connectivity: We observed significant activation of the nucleus reuniens (Re) located in the thalamus, which receives substantial afferent input from limbic structures and serves as a connector linking the hippocampus to the medial prefrontal cortex (Hauer et al., 2021). While our functional categorization offers a simplified overview, it’s worth noting that each activated brain region may have multifaceted roles beyond this classification. In summary, our study identified extensive activation of cortical and subcortical nuclei during KET anesthesia, encompassing regions related to arousal, pain modulation, neuroendocrine regulation, movement, and connectivity.

c-Fos expression in distinct brain regions after exposure to KET.

Earlier version: (A) Representative immunohistochemical staining of MOB, AON, ORB, MPO, ACA, MO, TRS, PL, ILA, DP, LS, PVT, SO, PVH, RE, VISC, AI, CLA, EPd, PIR, COA, AUD, TEa, ECT, PERI, CeA, SS, DG, STN, RSP, APN, LAT, EW, DR, PAG, SLD, PB, TRN, NI, LC, NTS, and NI c-Fos+ cells from the indicated mice. Scale bar, 200 µm. (B) Representation of cell counts in brain regions exhibiting significant statistical differences. Data are shown as mean ± SEM. P values < 0.05. (C) Schematic cross-section of the mouse brain showing activated brain regions by KET. Different colors indicate distinct functional nuclei. The red nuclei are associated with the regulation of sleep-wakefulness, the blue-green nuclei are linked to analgesia, the yellow nuclei are associated with motor function, and the white nuclei are a composite of various functional nuclei.

Revised version: (A) Representative immunohistochemical staining of MOB, AON, ORB, MPO, ACA, MO, TRS, PL, ILA, DP, LS, PVT, SO, PVH, RE, VISC, AI, CLA, EPd, PIR, COA, AUD, TEa, ECT, PERI, CeA, SS, DG, STN, RSP, APN, LAT, EW, DR, PAG, SLD, PB, TRN, NI, LC, NTS, and NI c-Fos+ cells from the indicated mice. Scale bar, 200 µm. (B) Representation of c-Fos density in brain regions exhibiting significant statistical differences. Data are shown as mean ± SEM. P values < 0.05. (C) Schematic cross-section of the mouse brain showing activated brain regions by KET. Different colors indicate distinct functional nuclei. The red nuclei are associated with the regulation of sleep-wakefulness, the blue-green nuclei are linked to analgesia, the yellow nuclei are associated with motor function, and the white nuclei are a composite of various functional nuclei. Figure 4—figure supplement 1. c-Fos expression in specific brain regions following normal saline administration. (A) Representative immunohistochemical staining of c-Fos+ cells in MOB, AON, ORB, MPO, ACA, MO, TRS, PL, ILA, DP, LS, PVT, SO, PVH, RE, VISC, AI, CLA, EPd, PIR, COA, AUD, TEa, ECT, PERI, CeA, SS, DG, STN, RSP, APN, LAT, EW, DR, PAG, SLD, PB, TRN, NI, LC, and NTS. Scale bar represents 200 µm.

Identification of brain regions activated by ISO

Earlier version: In our study, we aimed to identify brain regions activated by ISO and compare these with those activated by KET. We began by summarizing previously reported ISO activated nuclei, including PIR, LSd/LSv, CeA in the cortex and striatum, and VLPO, MnPO, EW, NTS, LC, ventral group of the dorsal thalamus (VeN), and area postrema (AP) in the hypothalamus and midbrain (Supplementary Table 3). We subsequently conducted a comprehensive MATLAB-based analysis of c-Fos expression throughout the entire brain, uncovering previously undetected activated nuclei (Figure 5A, Supplementary Figure 2). Newly identified activated nuclei in the CTX and the CNU included PL/ILA and ENT, aPVT in the thalamus, TU, ARH, PVi, and PVH in the hypothalamus, and PB in the hindbrain. All nuclei activated by ISO in this study were functionally classified and depicted in Figure 5C. Our results confirmed the activation of several brain regions involved in arousal-related nuclei, such as the prelimbic area (PL) and infralimbic areas (ILA), and the paraventricular nucleus (PVT) [22, 23]. Additionally, we observed activation in previously reported analgesia-related nuclei, including CeA and LC, as well as the parabrachial nucleus (PB) [18, 30, 35]. We also found activation in neuroendocrine function-related nuclei of the hypothalamus, such as TU, PVi, ARH, PVH, and SON. Moreover, we identified activations related to ISO induced side effects, such as in the piriform cortex (PIR) [36] and ENT [37], which may be stimulated by ISO odor, and the solitary tract nucleus (NTS), potentially responsible for ISO induced vomiting [38]. The only activated nucleus in the midbrain was the Edinger-Westphal nucleus (EW). Recent research has found that sevoflurane activates EW and is involved in sleep induction and maintenance of anesthesia, suggesting its crucial role in general anesthesia [39]. By comparing the ISO and KET induced c-Fos expression, we summarized the brain regions activated by both anesthetics in Figure 5D. Despite variations in molecular targets, the coactivation of regions such as PL/ILA, aPVT, CeA, PVH, SON, EW, PB, LC, and NTS by both ISO and KET suggests an overlapping neuronal circuitry that influences sleep-wake regulation, analgesia, and neuroendocrine functions. This shared neural circuitry may potentially offer a common mechanism across the two anesthetics for the maintenance of general anesthesia.

Revised version: To identify brain regions activated by ISO and compare these with those activated by KET. We began by summarizing previously reported ISO activated nuclei, including PIR, LSd/LSv, CeA in the cortex and striatum, and VLPO, MnPO, EW, NTS, LC, ventral group of the dorsal thalamus (VeN), and area postrema (AP) in the hypothalamus and midbrain (Figure 5—source data 1). We subsequently conducted a comprehensive MATLAB-based analysis of c-Fos expression throughout the entire brain, uncovering previously undetected activated nuclei (Figure 5A, Figure 5—figure supplement 1). Newly identified activated nuclei in the CTX and the CNU included PL/ILA and ENT, aPVT in the thalamus, TU, ARH, PVi, and PVH in the hypothalamus, and PB in the hindbrain. All nuclei activated by ISO in this study were functionally classified and depicted in Figure 5C. Our results confirmed the activation of several brain regions involved in arousal-related nuclei, such as the prelimbic area (PL) and infralimbic areas (ILA), and the paraventricular nucleus (PVT) (Mashour et al., 2022; Gao et al., 2020). Additionally, we observed activation in previously reported analgesia-related nuclei, including CeA and LC, as well as the parabrachial nucleus (PB) (Jiang-Xie et al., 2019; Llorca-Torralba et al., 2021; Deng et al., 2020). We also found activation in neuroendocrine function-related nuclei of the hypothalamus, such as TU, PVi, ARH, PVH, and SO. Moreover, we identified activations related to ISO induced side effects, such as in the piriform cortex (PIR) and ENT (Bekkers and Suzuki, 2013; Xu and Wilson, 2012), which may be stimulated by ISO odor, and the solitary tract nucleus (NTS), potentially responsible for ISO induced vomiting (Gupta et al., 2017). The only activated nucleus in the midbrain was the Edinger-Westphal nucleus (EW). Recent research has found that sevoflurane activates EW and is involved in sleep induction and maintenance of anesthesia, suggesting its crucial role in general anesthesia (Yi et al., 2023). By comparing the ISO and KET induced c-Fos expression, we summarized the brain regions activated by both anesthetics in Figure 5D. Despite variations in molecular targets, the coactivation of regions such as PL/ILA, aPVT, CeA, PVH, SO, EW, PB, LC, and NTS by both ISO and KET suggests an overlapping neuronal circuitry that influences sleep-wake regulation, analgesia, and neuroendocrine functions. This shared neural circuitry may potentially offer a common mechanism across the two anesthetics for the maintenance of general anesthesia.

c-Fos expression in distinct brain regions after exposure to ISO.

Earlier version: (A) Representative of brain regions with statistical differences c-Fos+ cells between the ISO group and home cage mice. Scale bar, 200 µm. (B) Representation of cell counts in brain regions exhibiting significant statistical differences between home cage and ISO. Data are shown as mean ± SEM. P values < 0.05. (C) Schematic cross-section of the mouse brain showing activated brain regions by ISO. Different colors indicate various functionally relevant nuclei. Red signifies nuclei involved in sleep-wake regulation, blue-green in pain management, blue in neuroendocrine function, pink in side-effect management, and white denotes nuclei exhibiting mixed functionalities. (D) The Venn diagram shows brain regions that are co-activated by ISO and KET and differentially activated brain regions.

Revised version: (A) Representative of brain regions with statistical differences c-Fos+ cells between the ISO group and home cage mice. Scale bar, 200 µm. (B) Representation of cell counts in brain regions exhibiting significant statistical differences between home cage and ISO. Data are shown as mean ± SEM. *P<0.05, **P<0.01, ***P<0.001. (C) Schematic cross-section of the mouse brain showing activated brain regions by ISO. Different colors indicate various functionally relevant nuclei. Red signifies nuclei involved in sleep-wake regulation, blue-green in pain management, blue in neuroendocrine function, pink in side-effect management, and white denotes nuclei exhibiting mixed functionalities. (D) The Venn diagram shows brain regions that are co-activated by ISO and KET and differentially activated brain regions. Figure 5—figure supplement 1. c-Fos expression in home cage group. (A) Representative immunohistochemical staining of c-Fos+ cells in PL, ILA, LSc, LSr, PIR, BST, VLPO, PVH, aPVT, SO, CeA, TU, PVi, ARH, EW, ENT, PB, LC, and NTS c-Fos+ cells from the indicated mice. Scale bar, 200 µm.

Network generation and Hub identification

Earlier version: Previous research has established that general anesthesia is mediated by various brain regions [7, 16–18]. c-Fos expression serves as an indicator of neuronal activity, providing a single index of activation per region per animal. By examining the covariance of c-Fos expression across animals within each group, we can infer interactions between brain regions and identify functional networks engaged during general anesthesia [40]. Consequently, we initially incorporated a total of 63 brain regions, including nuclei that were significantly activated by ISO and KET (as indicated by positive findings in Figures 4 and 5) and those previously reported to be associated with the mechanisms of general anesthesia and sleep-wakefulness [41]. Highly correlated activities across different brain regions are presumed to constitute components of a functional network, reflecting the complex interplay of brain areas during general anesthesia. Using c-Fos expression as a marker for neuronal activity and calculated a comprehensive set of interregional correlations for four groups. The matrices exhibited interregional correlations for the number of c-Fos-positive cells in each condition (Figure 6A). Network graphs were generated by extracting Pearson’s coefficients > 0.82, as well as significant positive correlations (P < 0.05), from these matrices to construct functional brain networks (Figure 6B). Compared to the control group, isoflurane (ISO) slightly decreased interconnections between regions (network density: 0.13 vs. 0.10; edges: 267 vs. 198), with no significant difference in mean interregional correlation coefficients between the ISO group and the home cage group (Fisher Z = −0.018, P = 0.98). Conversely, ketamine (KET) significantly increased the network’s connectivity density (network density: 0.13 vs. 0.47; edges: 265 vs. 1008) and showed a significant increase in mean interregional correlation coefficients compared to the saline group (Fisher Z = 3.54, P < 0.001) (Figure 6C). These findings suggest that ISO may exert a mild inhibitory effect on functional network connectivity, whereas KET appears to enhance interregional correlation following administration (Figure 6D).

Revised version: Previous research has established that general anesthesia is mediated by various brain regions (Moore et al., 2012; Gelegen et al., 2018; Hua et al., 2020; Jiang-Xie et al., 2019). C-Fos expression serves as an indicator of neuronal activity, providing a single index of activation per region per animal. By examining the covariance of c-Fos expression across animals within each group, we can infer interactions between brain regions and identify functional networks engaged during general anesthesia (Wheeler et al., 2013). We initially incorporated a total of 63 brain regions, including nuclei that were significantly activated by ISO and KET (as indicated by positive findings in Figures 4 and 5) and those previously reported to be associated with the mechanisms of general anesthesia and sleep-wakefulness (Leung et al., 2014). We first calculated a comprehensive set of interregional correlations for four groups (Figure 6A). Correlation analysis revealed similar network connectivity and densities between the ISO and home cage groups (mean interregional correlations: Fisher Z = −0.018, P = 0.98; network densities: 0.10 vs. 0.09). In contrast, KET significantly increased network density (0.46 vs. 0.06) and mean interregional correlations compared to the saline group (Fisher Z = 3.54, P < 0.001; Figure 6—figure supplement 1D), indicating enhanced interregional correlation with KET administration. Additionally, we segmented the functional network into modules based on the hierarchical clustering of correlation coefficients (Figure 6—figure supplement 1A). We chose to cut the tree at a value of 0.7 for the relatively high value of Silhouette scores, resulting in two modules for the home cage and KET groups, and four for the saline and ISO groups (Figure 6—figure supplement 1B and C). Networks were formed using Pearson’s coefficients over 0.82 and significant correlations (P < 0.05), leaving a single module for KET and three modules for ISO (Figure 6B).

Generation of anesthetics-induced networks and identification of hub regions.

Added in the revised version: (A) Heatmaps display the correlations of log c-Fos densities within brain regions (CTX, CNU, TH, HY, MB, and HB) for various states (home cage, ISO, saline, KET). Correlations are color-coded according to Pearson’s coefficients. The brain regions within each anatomical category are organized by hierarchical clustering of their correlation coefficients. Full names and expression levels for each brain region are detailed in Supplementary Figure 3. (B) Network diagrams illustrate significant positive correlations (P < 0.05) between regions, with Pearson’s r exceeding 0.82. Edge thickness indicates correlation magnitude, and node size reflects the number of connections (degree). Node color denotes betweenness centrality, with a spectrum ranging from dark blue (lowest) to dark red (highest). The networks are organized into modules consistent with the clustering depicted in Figure 6—figure supplement 1A. Figure 6—figure supplement 1. Hierarchical clustering across different conditions. (A) Heatmaps depict the hierarchical clustering of brain regions within the home cage, ISO, saline, and KET groups, using Euclidean distance matrices calculated from correlation coefficients. Each brain region is represented by its abbreviation, with full names and expression levels detailed in Figure 3—figure supplement 1. Modules are demarcated within each dendrogram at a cut-off threshold of 0.7. (B) Silhouette scores are plotted against the dendrogram tree height ratio for each condition, with optimal cluster definition indicated by a dashed line at a 0.7 ratio. (C) The number of clusters formed at different cutoff levels. At a ratio of 0.7, ISO and saline treatments result in three clusters, whereas home cage and KET conditions yield two clusters. (D) The mean Pearson’s correlation coefficient (r) was computed from interregional correlations displayed in Figure 6A. Data were analyzed using one-way ANOVA with Tukey’s post hoc test, ***P < 0.001. Figure 6—figure supplement 2. Hub region characterization across different conditions: home cage (A), ISO (B), saline (C), and KET (D) treatments. Brain regions are sorted by degree, betweenness centrality, and eigenvector centrality, with each metric presented in separate bar graphs. Bars to the left of the dashed line indicate the top 20% of regions by rank, highlighting the most central nodes within the network. Red

Removed in the revised version: (A) Matrices showing interregional correlations for c-Fos expression at the home cage, ISO, saline, and KET. Colors reflect correlation strength based on Pearson’s r (color bar, right). Axes are numbered and correspond to brain regions listed in Supplementary Table 5. (B) Network graphs were generated by significant positive correlations (P < 0.05), as well as Pearson’s r > 0.82. The widths of the edges are proportional to the strength of the correlations and the size of the nodes is proportional to the degree of the node. Colors represent the major brain division (red, cerebral nuclei; purple, hypothalamus; orange, midbrain; brown, thalamus; blue, thalamus; green, the cerebral cortex. Network densities were noted in the left. (C) Mean r was calculated from interregional correlation coefficients. Data were analyzed using one-way ANOVA, followed by Tukey’s post hoc test. ***P < 0.001.

Removed in the revised version: Hubs are nodes that occupy critical central positions within the network, enabling the network to function properly. To determine the brain regions that serve as the hubs for ISO and KET induced functional networks, we calculated each node’s degree (the number of links that connect it) and betweenness centrality (the fraction of all shortest paths in the network that contain a given node). Nodes with high degree and betweenness centrality values typically exhibit many connections. We ranked each node according to its degree and betweenness centrality and extracted nodes with a rank greater than the 80th percentile in its network (Figure 7). Additionally, we segmented the functional network into non-overlapping modules using a spectral community detection algorithm and calculated the within-modal degree Z-score (the degree of nodes within a module, which indicates within-module connectivity) as well as participation coefficients (the distribution of the edges of a node between other modules, which indicates within-module connectivity) [42]. Nodes with relatively high values for both parameters were considered relatively significant within and outside the network module. We classified nodes with participation coefficients > 0.4 and Z scores > 1 as connector hubs [43]. The LC, exhibiting elevated degree and betweenness centrality as well as relatively higher Z-scores and participation coefficients, may play a pivotal role in mediating isoflurane-induced general anesthesia, given its known involvement in arousal, attention, and analgesia. In contrast, the somatosensory cortex (SS) functions as a connector hub in the KET group, indicating its integrative and coordinating role in ketamine-induced dissociative anesthesia (Figure 7D). The LSc served as the central hub in the saline group (Figure 7C). No significant findings were observed in the within-modal degree Z-score and participation coefficient analyses for the home cage group. However, the analysis did reveal a relatively high degree and betweenness centrality of APNs (Figure 7A), suggesting that the within-module connections are relatively independent, or that the pathways of inter-module information transmission do not depend on specific connector nodes.

Removed in the revised version: Identification of hub regions in the home cage (A), ISO (B), NS (C), and KET (D) groups.

Removed in the revised version: Each brain region is ranked in descending order based on degree and betweenness centrality, respectively. The regions to the left of the dotted line represent the top 20% in our network. Black columns highlight brain regions that rank similarly in both degree and betweenness centrality. The within-module Z-scores and participation coefficients were calculated for each brain region in the networks. Brain regions characterized by a participation coefficient greater than 0.4 and a within-module degree Z-score greater than 1 are defined as hubs and are denoted by red dots.

Earlier version: The possible framework for ketamine and isoflurane induced unconsciousness.

Revised version: The possible framework for KET and ISO-induced unconsciousness.

Earlier version: The distinct pathways of ketamine and isoflurane induced unconsciousness can be explained by two contrasting mechanisms. The “top-down” process attributes KET’s effect to widespread cortical activation (represented in yellow), with the somatosensory cortex (SS) acting as the central node in the functional network (depicted in blue) and relative inhibition of hypothalamic sleep-promoting nuclei. Conversely, the “bottom-up” approach posits that isoflurane induced unconsciousness arises from the activation of hypothalamic sleep-promoting regions (indicated in yellow) and relative inhibition of cortical and thalamic nuclei, with the locus coeruleus (LC) serving as the hub node in the isoflurane induced functional network. Nuclei activated by both anesthetics are shown in green. Adapted from [3, 4, 52]. PL, prelimbic area; ILA, infralimbic areas; SON, supraoptic nucleus; PVH, paraventricular hypothalamic nucleus; TU, tuberal nucleus; LC, locus coeruleus; SS, somatosensory cortex; CTX: cortex; TH: thalamus; HY, hypothalamus; MB; midbrain; HB, hindbrain.

Revised version: The distinct pathways of KET and ISO-induced unconsciousness can be explained by two contrasting mechanisms. The ‘top-down’ process attributes KET’s effect to widespread cortical activation (represented in yellow), with the temporal association areas (TEa) acting as the central node in the functional network (depicted in blue). Conversely, the ‘bottom-up’ approach suggests that isoflurane-induced unconsciousness stems from the activation of certain hypothalamic regions (highlighted in yellow), with the locus coeruleus (LC) acting as the hub node within the isoflurane-induced functional network. Nuclei activated by both anesthetics are shown in green. Adapted from (Mashour, 2014; Mashour and Hudetz, 2017; Reimann and Niendorf, 2020). PL, prelimbic area; ILA, infralimbic areas; SO, supraoptic nucleus; PVH, paraventricular hypothalamic nucleus; LC, locus coeruleus; SS, somatosensory cortex; CTX: cortex; TH: thalamus; HY, hypothalamus; MB; midbrain; HB, hindbrain.

Added in the revised version: Hubs are nodes that occupy critical central positions within the network, enabling the network to function properly. Due to the singular module structure of the KET network and the sparsity of intermodular connections in the home cage and saline networks, the assessment of network hub nodes did not employ within-module degree Z-score and participation coefficients, as these measures predominantly underscore the importance of connections within and between modules (Kimbrough et al., 2020). The analysis evaluated nodal importance using degree (number of direct connections), betweenness centrality (proportion of shortest network paths passing through a node), and eigenvector centrality (influence based on the centrality of connected nodes), identifying brain regions that scored highly in all three as hub nodes (Figure 6—figure supplement 2). The locus coeruleus (LC), characterized by its high degree and eigenvector centrality, as well as its notably higher betweenness centrality, underscores its pivotal role in propagating and integrating neural signals within the ISO-influenced network. In the KET group, the temporal association areas (TEa) serve as a hub, emphasizing their integrative function in the manifestation of ketamine-induced dissociative anesthesia. Meanwhile, the PAG and the lateral septal nucleus, rostral part (LSr), function as central hubs in the home cage and saline groups, respectively.

Discussion

Added in the revised version: In this study, a comparative analysis was conducted on the effects of two general anesthetics, ISO and KET, on c-Fos expression throughout the brain. Hierarchical cluster analysis enabled a detailed examination of individual brain region responses to each anesthetic. We found that KET primarily activates the cerebral cortex, indicative of a top-down mechanism, whereas ISO chiefly stimulates subcortical areas, particularly the hypothalamus, aligning with a bottom-up approach (Mashour, 2014; Mashour and Hudetz, 2017). Further analysis based on c-Fos expression revealed the TEa as a key hub in the KET-induced network and the LC in the ISO network, underscoring their significant roles in mediating the effects of KET and ISO, respectively.

Added in the revised version: Our hierarchical clustering method effectively identified brain regions impacted by ISO and KET, aligning with multiple comparison results as shown in Figures 4 and 5. This approach facilitated the efficient grouping of regions with similar expression patterns, enabling a comprehensive analysis of brain-wide c-Fos expression differences. Notably, certain brain regions such as SO, VLPO, TU, and CeA were prominent under ISO treatment, with SO (Jiang-Xie et al., 2019), VLPO (Moore et al., 2012), and CeA (Hua et al., 2020) previously reported to be involved in mechanisms of general anesthesia. The LC was identified as a highly connected hub, underscoring its critical role in the brainstem during ISO-induced general anesthesia. The LC is known for its diverse functions, including arousal, pain modulation, attention, stress response, and neuroprotection. Studies indicate that chemical activation of the LC enhances whole-brain functional connectivity, largely due to its role as the primary source of norepinephrine (NE) and its extensive influence throughout the brain (Zerbi et al., 2019). The significant activation of the LC and its pivotal position in the functional network underlying ISO-induced unconsciousness indicate its crucial role in maintaining and integrating the network. This finding highlights the LC’s involvement in the bottom-up paradigm of ISO-induced unconsciousness.

Added in the revised version: KET induces extensive activation across the brain, with a notable impact on cortical and cerebral nuclei, as indicated in Figures 3 and 4. This observation aligns with the proposal by Mashour et al. that anesthesia-induced unconsciousness includes both the modulation of bottom-up and top-down neural processing (Mashour, 2014; Mashour and Hudetz, 2017). In this framework, anesthetics like KET diminish consciousness by interfering with cortical and thalamocortical circuits, which are crucial for neural information integration. Our study found that KET administration significantly activates cortical and subcortical arousal-promoting nuclei. Graph theoretical analysis identified the TEa as a central hub in the KET-induced functional network. Given TEa’s critical role in integrating multisensory information, particularly auditory signals, and its involvement in complex behaviors such as maternal responses as evidenced by its role in processing pup ultrasonic vocalizations, our findings suggest that KET may influence the connectivity of TEa with other brain regions, thus potentially altering the overall dynamics of the functional brain network (Tasaka et al., 2020). This further supports the significance of cortical areas during KET anesthesia.

Removed in the revised version: In this study, we conducted a comparative analysis of the effects of two general anesthetics, isoflurane (ISO) and ketamine (KET), on c-Fos expression throughout the brain. By employing principal component analysis (PCA), we were able to thoroughly examine the responses of individual brain regions to each anesthetic agent. Our findings reveal that KET dominantly activates the cerebral cortex yet suppresses subcortical regions, reflecting a top down mechanism of action, while ISO predominantly stimulates subcortical brain regions with relative cortical inhibition, substantiating its bottom up mechanism of action [3, 4]. Further functional analysis of brain networks, based on c-Fos expression, identified the somatosensory cortex (SS) and the locus coeruleus (LC) as central nodes for KET and ISO, respectively, highlighting the crucial roles of LC and SS under ISO and KET induced unconsciousness.

Removed in the revised version: Our results demonstrate that ISO activates the sleep-promoting VLPO nucleus, the aPVT-infralimbic loop [23], and broadly inhibits the cortex, supporting a bottom-up mechanism of ISO induced unconsciousness. Nonetheless, our findings also reveal the activation of arousal related nuclei, such as PB and LC, which implies that the influence of isoflurane on consciousness may not solely rely on suppressing arousal centers, but rather through a more intricate relationship than formerly recognized. Identifying cell types and their dynamic changes during anesthesia will be crucial for clarifying their role in ISO induced unconsciousness.

Removed in the revised version: While c-Fos-based functional network analysis offers lower temporal resolution, its single-cell resolution across the entire brain allows for the inclusion of midbrain and hindbrain regions, supplementing previous fMRI analyses. Our observations reveal that ISO mildly inhibits network density, and through graph theoretical analysis, we identify the LC as a highly connected hub, highlighting the critical role of the brainstem in ISO induced general anesthesia. The LC performs a wide range of functions in mice, including arousal, pain modulation, attention, stress response, and neuroprotection. Studies have shown that chemical activation of the LC increases whole-brain functional connectivity, attributed to its role as the primary source of norepinephrine (NE) and its extensive influence on nearly the entire brain [44]. The significant LC activation and its central position within the functional network underlying ISO induced unconsciousness suggest that the LC plays a crucial part in maintaining and integrating the entire unconsciousness functional network, emphasizing the involvement of LC in the bottom-up paradigm of ISO induced unconsciousness.

Removed in the revised version: Mashour et al. proposed that anesthesia-induced unconsciousness encompasses not only the modulation of lower-level brain activity but also top down neural processing [3, 4]. Within this top down framework, anesthetics diminish consciousness by interfering with cortical and thalamocortical circuits responsible for neural information integration. Our study discovered that KET administration substantially activated cortical and subcortical arousal-promoting nuclei while concurrently causing relative thalamic suppression, with only the RE and TRS exhibiting activation. This suggests that thalamic inhibition may lead to a reduction in thalamocortical communication, which is characterized by the inability to perceive the external environment and results in disconnection from reality. Graph theoretical analysis also identified the somatosensory cortex (SS) as the hub node of the KET induced functional network. As a critical cortical area, SS is responsible for sensory processing, motor control, and cognitive functions [45]. Previous studies have demonstrated that local KET administration to SS recapitulates the effects of systemic KET on both the switch in pyramidal cell activity and dissociative-like behavior, implying that SS may serve as a key target for KET induced dissociation [46]. Our findings that SS acts as a hub node suggest that KET may modulate brain network function by influencing connectivity between SS and other brain regions, thereby affecting the behavior and cognitive states of mice. This further supports the significance of cortical areas during KET anesthesia.

Earlier version: Identifying shared neural features between KET and ISO is essential for understanding anesthetic-induced unconsciousness. The coactivation of sleep-wake regulation-related regions, such as PL/ILA and aPVT, along with analgesia-related nuclei like CeA, PB, and LC, suggests a shared mechanism for sleep-wake regulation and the common pathways for pain relief. This observation provides valuable insights into the fundamental mechanisms of anesthesia-induced hypnosis and analgesia. Additionally, the coactivation of neuroendocrine-related nuclei, including PVH and SO in the hypothalamus, raises questions about the potential influence of anesthetics on hormonal release and homeostatic regulation. Other coactivated nuclei, such as EW and NTS, warrant further investigation of their roles in anesthesia. In summary, the coactivated nuclei imply a potential shared neuronal circuitry for general anesthesia, encompassing common features like unconsciousness, analgesia, and autonomic regulation, regardless of the specific molecular targets of each drug. Future research could examine coactivated brain regions by the two anesthetics or manipulate identified hub nodes to further understand the mechanisms of general anesthesia. In summary, our study reveals distinct and shared neural mechanisms underlying isoflurane and ketamine anesthesia using c-Fos staining and network analysis. Our findings support “top-down” and “bottom-up” paradigms, and the identification of hub nodes and coactivated brain regions suggests shared neurocircuitry for general anesthesia, providing insights into the mechanisms underlying anesthetic-induced unconsciousness and analgesia.

Revised version: Identifying shared neural features between KET and ISO is crucial for understanding anesthetic-induced unconsciousness. The coactivation of brain regions, such as PL/ILA and aPVT, alongside analgesia-related nuclei like CeA, PB, and LC, suggests a shared mechanism for thalamocortical regulation and common pathways for pain relief. This observation provides valuable insights into the fundamental mechanisms of anesthesia-induced hypnosis and analgesia. Additionally, the coactivation of neuroendocrine-related nuclei, particularly PVH and SO in the hypothalamus, raises questions about the potential influence of anesthetics on hormonal release and homeostatic regulation. Further investigation is warranted for other coactivated nuclei, such as EW and NTS, to understand their roles in anesthesia. The coactivated nuclei suggest a potential shared neuronal circuitry for general anesthesia, encompassing common features like unconsciousness, analgesia, and autonomic regulation. Future research could focus on examining coactivated brain regions by the two anesthetics or manipulating identified hub nodes to delve deeper into the mechanisms of general anesthesia.

Added in the revised version: Our findings indicate that c-Fos expression in the KET group is significantly elevated compared to the ISO group, and the saline group exhibits notably higher c-Fos expression than the home cage group, as seen in Figure 3—figure supplement 1. Intraperitoneal saline injections in the saline group, despite pre-experiment acclimation with handling and injections for four days, may still evoke pain and stress responses in mice. Subtle yet measurable variations in brain states between the home cage and saline groups were observed, characterized by changes in normalized EEG delta/theta power (home cage: 0.05±0.09; saline: −0.03±0.11) and EMG power (home cage: −0.37±0.34; saline: 0.04±0.13), as shown in Figure 1—figure supplement 1. These changes suggest a relative increase in overall brain activity in the saline group compared to the home cage group, potentially contributing to the higher c-Fos expression. Although the difference in EEG power between the ISO group and the home cage control was not significant, the increase in EEG power observed in the ISO group was similar to that of KET (0.47 ± 0.07 vs 0.59 ± 0.10), suggesting that both agents may induce loss of consciousness in mice. Regarding EMG power, ISO showed a significant decrease in EMG power compared to its control group. In contrast, the KET group showed a lesser reduction in EMG power (ISO: −1.815± 0.10; KET: −0.96 ± 0.21), which may partly explain the higher overall c-Fos expression levels in the KET group. This is consistent with previous studies where ketamine doses up to 150 mg/kg increase delta power while eliciting a wakefulness-like pattern of c-Fos expression across the brain (Lu et al., 2008). Furthermore, the observed differences in c-Fos expression may arise in part from the dosages, routes of administration, and their distinct pharmacokinetic profiles. This variation is compounded by the lack of detailed physiological monitoring, such as blood pressure, heart rate, and respiration, affecting our ability to precisely assess anesthesia depth. Future studies incorporating comprehensive physiological monitoring and controlled dosing regimens are essential to further elucidate these relationships and refine our understanding of the effects of anesthetics on brain activity.

Added in the revised version: Methods and Materials

Removed in the revised version: Methods

Animals

All animal experiments were conducted in accordance with the National Institutes of Health guidelines and were approved by the Chinese Academy of Sciences’ Institute of Neuroscience. Adult male wild-type (WT) mice (C57BL/6J) (8-10 weeks old, weight from 22 to 26 g) were purchased from institute-approved vendors (LingChang Experiment Animal Co., China). Mice were individually housed and maintained on a 12 h:12 h light/dark cycle (lights on at 07:00 a.m. and off at 07:00 p.m.) with food and water available ad libitum.

Drug administration

Earlier version: All experiments occurred between 13:00-14:30 (ZT6–ZT7.5). We adapted mice to handling and the anesthesia chamber (10×15×15 cm) for several days to minimize experimental confound-induced c-Fos induction. Adult male mice were handled for the KET group for 10 min per day with normal saline (NS) injected intraperitoneally (i.p.) for three consecutive days at 13:00. On day five, a randomly chosen mouse received an injection of Ketamine (Gutian Medicine, H35020148), and the control groups (n=8) received the same volume of saline. ISO group (n=6) mice were handled and inhaled 1.5% isoflurane (RWD Life Science, 1903715) at 13:00 on day four in the chamber. Meanwhile, the control groups (n=6) were left undisturbed in their home cages prior to sampling. We confirmed the loss of righting reflex at 5 min after anesthetics exposure. For 90 min after KET injection or ISO inhalation, mice were deeply anesthetized with 5% ISO and transcardially perfused with 35 ml 0.1 M phosphate-buffered saline (PBS) followed by 35 ml 4% paraformaldehyde (PFA). The brains were then removed and postfixed overnight with 4% PFA. Following fixation, the brains were dehydrated for 48 hours with 30% sucrose (wt/vol) in PBS. Coronal sections (50 µm) of the whole brain were cut using a cryostat (HM525 NX, Thermo Scientific) after being embedded with OCT compound (NEG-50, Thermo Scientific) and freezing.

Revised version: All experiments were conducted between 13:00-14:30 (ZT6–ZT7.5). We adapted mice to handling and the anesthesia chamber (10×15×15 cm) for four days to minimize experimental confound-induced c-Fos induction. For KET administration, adult male mice were handled for 10 minutes per day, and normal saline (NS) was injected intraperitoneally (i.p.) for four consecutive days at 13:00. On day five, a randomly chosen mouse received an injection of ketamine (Gutian Medicine, H35020148) (n=6), and the control groups (n=8) received the same volume of saline. ISO group (n=6) mice were handled and inhaled 1.5% isoflurane (RWD Life Science, 1903715) at 13:00 on day four in the chamber. Meanwhile, the control groups (n=6) were left undisturbed in their home cages prior to sampling. We confirmed the loss of righting reflex at 5 min after anesthetics exposure. For 90 min after KET injection or ISO inhalation, mice were deeply anesthetized with 5% ISO and transcardially perfused with 35 ml 0.1 M phosphate-buffered saline (PBS) followed by 35 ml 4% paraformaldehyde (PFA). The brains were then removed and postfixed overnight with 4% PFA. Following fixation, the brains were dehydrated for 48 hours with 30% sucrose (wt/vol) in PBS. Coronal sections (50 µm) of the whole brain were cut using a cryostat (HM525 NX, Thermo Scientific) after being embedded with OCT compound (NEG-50, Thermo Scientific) and freezing.

Anesthesia depth measurement

Earlier version: EEG and EMG recordings were conducted as described previously[20]. Specifically, a stainless steel electrode was positioned over the frontal cortex for EEG measurements, while an insulated EMG electrode was inserted into the neck muscles. Additionally, a reference electrode was placed above the cerebellum. The EEG/EMG data were captured using TDT system-3 amplifiers, specifically RZ5 + RA16PA and RZ2 + PZ5 configurations. To measure the depth of anesthesia using EEG/EMG signals, a fast Fourier transform spectral analysis was employed, featuring a frequency resolution of 0.18 Hz. Data was noted every 5 seconds through a MATLAB tool and subsequently verified manually by experts. The assessment was based on EMG power and the ratio of EEG’s δ power (0.5-4Hz) to θ power (6-10Hz). EEG and EMG power values within 30 minutes post-administration were normalized to a 5-minute pre-administration baseline.

Revised version: EEG and EMG recordings were conducted as previously described (Luo et al., 2023). Specifically, a stainless steel electrode was positioned over the frontal cortex for EEG measurements, while an insulated EMG electrode was inserted into the neck muscles. Additionally, a reference electrode was placed above the cerebellum. The EEG/EMG data were captured using TDT system-3 amplifiers, specifically RZ5 + RA16PA and RZ2 + PZ5 configurations. To measure the depth of anesthesia using EEG/EMG signals, a fast Fourier transform spectral analysis was employed, featuring a frequency resolution of 0.18 Hz. Data was noted every 5 seconds through a MATLAB tool and subsequently verified manually by experts. The assessment was based on EMG power and the ratio of EEG’s δ power (0.5-4Hz) to θ power (6-10Hz). EEG and EMG power values within 30 minutes post-administration were normalized to a 5-minute pre-administration baseline.

Immunohistochemistry

One out of every three brain slices (100 µm intervals) of each whole brain was washed three times with 0.1 M phosphate-buffered saline (PBS) for 10 min and permeabilized for 30 minutes at room temperature with 0.3% Triton X-100 in PBS (PBST). Slices were incubated for 2 hours at room temperature with 2% normal donkey serum (Sigma, G6767) in PBS overnight at 4°C with c-Fos primary antibodies (226003, Synaptic Systems; 1:500) diluted in PBS with 1% donkey serum. After three washes with PBST, slices were incubated with the Cy3 donkey anti-rabbit (711165152, Jackson; 1:200) secondary antibody for 2 hours at room temperature. Immunostained slices were mounted with VECTASHIELD mounting medium with DAPI and then scanned under a fluorescent microscope equipped with a 10× objective (VS120, Olympus) or a confocal microscope with a 20× objective (FV300, Olympus). Due to the limited scope of a single field of view, either 2 or 4 adjacent fields of view were stitched together to offer a comprehensive representation of specific brain regions, including the PL, ILA, LSc, LSr, SS, and VISC.

Quantification of c-Fos positive cells

Earlier version: The procedures used for c-Fos analysis were based on previous research [19]. A custom-written software package was employed for cellular quantification within brain images. The software consists of three components: atlas rotation, image registration, and signal detection.

Revised version: The procedures used for c-Fos analysis were based on previous research (Ma et al., 2021). A custom-written software package was employed for cellular quantification within brain images. The software consists of three components: atlas rotation, image registration, and signal detection.

Atlas rotation module

Utilizing the Allen Mouse Brain Atlas, this module allows rotations in both horizontal and vertical dimensions to align with mouse brain samples. To determine the appropriate rotation angles, we manually pinpointed anatomical landmarks corresponding to the samples. For the left-right axis rotation, we chose points from the CA3 pyramidal layer in both hemispheres and the posterior slice of the dentate gyrus, specifically the granule cell layer. For the dorsal-ventral axis rotation, we identified key anatomical landmarks. These include the initial connections of the anterior commissure and corpus callosum between the hemispheres, as well as ventral regions like the interpeduncular nucleus and the suprachiasmatic nucleus. After determining these rotation angles, we adjusted the reference atlas to match our samples.

Image registration module

This module uses a tool that aligns brain slice images with a reference atlas, facilitating the alignment of overall brain region distribution. Registration starts by matching the coronal plane of the sample sections to the atlas. After defining the external boundaries of the brain section, the system performs geometric transformations on the section to optimize its fit with the atlas.

Signal detection module

The detection module is designed to automatically label c-Fos+ cells. Post detection, a manual verification was performed on each digitized brain section image to ensure the accuracy and precision of the c-Fos+ cell markings.

Added in the revised version: Hierarchical clustering

Added in the revised version: Prior to hierarchical clustering in Figures 2 and 3, we calculated the relative c-Fos densities by dividing the c-Fos densities of each brain region in the experimental groups by their respective controls and then performed a log transformation on these values to obtain the log relative c-Fos densities. These log ratios, which normalize the data and reduce variance, indicate differential expression with a value of zero denoting no change compared to control. After normalizing the data, we performed hierarchical clustering by first computing the pairwise Euclidean distances among brain regions. Regions with the smallest distances, indicating high similarity, were grouped iteratively. Cluster boundaries were defined using complete linkage, ensuring homogeneity within clusters by considering the largest distance between members. In Figure 6, hierarchical clustering was performed within each of the CTX, CNU, TH, HY, MB, and HB regions based on the log-transformed c-Fos density correlations. In Figure 6—figure supplement 1, hierarchical clustering was performed based on the interregional log c-Fos density correlations to identify modules of coactivation within each treatment group, revealing underlying functional connectivity networks (Kimbrough et al., 2020). To determine the statistical significance of c-Fos expression differences, we computed the Z-score for each treatment condition—ISO and KET—by dividing the log relative c-Fos density by the standard error. Positive z-scores indicate higher values than control, and negative z-scores indicate lower values than control.

Removed in the revised version: Principal components analysis (PCA) of the activity patterns at whole brain

Removed in the revised version: We utilized arithmetic methods to compute a logarithm ratio of the c-Fos density in a specific brain region for PCA analysis. Specifically, the cell density of a single brain region (defined as the number of c-Fos positive cells in a brain region divided by the volume of that region) was divided by the global c-Fos cell density (calculated as the total number of c-Fos positive cells throughout the entire brain divided by the total brain volume). Subsequently, the logarithm of this ratio value was taken. The formula is as follows: Nr: The number of c-Fos+ cells in each brain region. Vr: The volume of each brain region. Principal components analysis was performed on the concatenated matrix containing data from four conditions, using singular value decomposition. We then selected the first two principal components (PCs), which accounted for 56.6% of the variance in Figure 2A and 44.5% in Figure 3A.

Network generation

Earlier version: To evaluate how functional connectivity changed under general anesthetics in WT mice, we extracted 63 brain regions from major brain subdivisions (cerebral cortex, cerebral nuclei, thalamus, hypothalamus, midbrain, and hindbrain) listed in Supplementary Table 2. Correlation matrices were generated by computing Pearson correlation coefficients from interregional c-Fos expression in the 63 regions. Mean correlations were calculated to assess changes in functional connectivity between these major subdivisions of the brain. Weighted undirected networks were constructed by considering correlations with Pearson’s r ≥0.82, corresponding to a one-tailed significance level of P<0.05 (uncorrected). The nodes in the networks represent brain regions, and the correlations that survived thresholding were considered connections. Theoretical graph analysis was performed using Brain Connectivity Toolbox (https://sites.google.com/site/bctnet/, version 2019-03-03) in MATLAB R2021 (he MathWorks Inc.) [47]. Network visualization was performed using Cytoscape (version 3.2.1) [48].

Revised version: To evaluate how functional connectivity changed under general anesthetics in WT mice, we extracted 63 brain regions from major brain subdivisions (cerebral cortex, cerebral nuclei, thalamus, hypothalamus, midbrain, and hindbrain) listed in Figure 6A. Correlation matrices were generated by computing Pearson correlation coefficients from interregional c-Fos expression in the 63 regions. Mean correlations were calculated to assess changes in functional connectivity between these major subdivisions of the brain. Weighted undirected networks were constructed by considering correlations with Pearson’s r ≥0.82, corresponding to a one-tailed significance level of P<0.05 (uncorrected). The nodes in the networks represent brain regions, and the correlations that survived thresholding were considered connections. Theoretical graph analysis was performed using Brain Connectivity Toolbox (https://sites.google.com/site/bctnet/, version 2019-03-03) in MATLAB R2021 (The MathWorks Inc.) (Rubinov and Sporns, 2010). Network visualization was performed using Cytoscape (version 3.2.1) (Shannon et al., 2003).

Hub identification

Added in the revised version: Network centrality was evaluated using degree, betweenness, and eigenvector centrality measures to identify potential hub regions (Guimera and Nunes Amaral, 2005). Degree counts the number of edges connected to a node, signifying its immediate influence. Betweenness centrality is gauged by the number of shortest paths passing through a node, indicating its role as a connector or ‘bridge’ within the network. Eigenvector centrality measures a node’s influence by the centrality of its connections, valuing nodes linked to well-connected neighbors. High eigenvector centrality indicates significant influence through these high-quality connections within the network.

Removed in the revised version: Network centrality was evaluated using degree, betweenness, within-module z-scores (representing within-module connectivity), and participation coefficient (indicating between-module connectivity). These measures were computed for all nodes to identify potential hub regions [49]. Modularity was assessed employing Newman’s spectral community detection algorithm [50]. Degree represents the number of edges connected to a node, while betweenness denotes the number of shortest paths traversing a given node. Nodes with elevated betweenness centrality are involved in numerous shortest paths.

Statistical analysis

Earlier version: Sample size was determined based on prior studies [13, 14]. Data are presented as mean ± SEM, and all statistical tests were two-sided. For the c-Fos cell percentages in Figure 2C, we employed the Mann-Whitney test with a false discovery rate correction, setting a Q-value at 0.05. In Figures 4 and 5, differences in c-Fos cell counts for both KET and ISO were evaluated using an independent t-test. P-values were adjusted using the procedure by Benjamini, Krieger, and Yekutieli with Q set at 0.05 [51]. Figure 6C and Supplementary Figure 1 utilized a one-way ANOVA followed by a post-hoc Tukey correction. Pearson correlation coefficients (R) were transformed into Z-scores using the Fisher’s Z transformation before computing group means and making statistical comparisons. All statistical analyses were conducted using GraphPad Prism 9.0 (GraphPad Software, USA) and MATLAB R2021 (Mathworks Inc.). A P-value of less than 0.05 was deemed statistically significant.

Revised version: Sample size was determined based on prior studies (Lu et al., 2008; Yatziv et al., 2020). Data are presented as mean ± SEM. In Figures 4 and 5, independent t-tests evaluating differences in c-Fos expression for KET and ISO treatments were corrected for multiple comparisons using the Benjamini, Krieger, and Yekutieli method, with a 5% false discovery rate (Q) (Benjamini and Hochberg, 1995). Pearson correlation coefficients (R) were transformed into Z-scores using Fisher’s Z transformation before computing group means and making statistical comparisons in Figure 6—figure supplement 1D. All statistical analyses were conducted using GraphPad Prism 9.0 (GraphPad Software, USA) and MATLAB R2021 (Mathworks Inc.).

Removed in the revised version: Supporting information